# **Problem Set 11**

## **Tasks**
- Assembling A and S matrices
- Calculating d
- Getting the direct and indirect GWP for both cases
- Answering the multiple choice questions

## **Passing requirement**
- **At least 80 points**

**_Please run the cell below before you start the assignment_**

```python
# Import required Python libraries
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns

# Folder for the files
files_folder = "files_PS11_V1_0/"

### In case formulas are entered in the csv files & correcting the datatypes
def convert_formulas(matrix):
    for row in matrix.index:
        for column in matrix.columns:
            if type(matrix.loc[row, column]) == str:
                if " " in matrix.loc[row, column]:
                    matrix.loc[row, column] = matrix.loc[row, column].lstrip(' ')
                if "=" in matrix.loc[row, column]:
                    matrix.loc[row, column] = matrix.loc[row, column].lstrip('=')
                matrix.loc[row, column] = pd.eval(matrix.loc[row, column])
    ### Get the data types right in the DataFrames
    s = matrix.select_dtypes(include='object').columns
    matrix[s] = matrix[s].astype("float64")
    return matrix
```

_**Disclaimer**: Mainstream LCA assessments normally use commercial LCI databases, such as Ecoinvent, for their background systems. These databases are, as the name indicates, commercial in nature and are as such not open access. We are therefore using an open-access mixed-unit version of EXIOBASE. This is an Input-Output database (IOT). So for your applications, you will in principle be using an IO-Hybrid LCA model rather than a conventional process-based LCA model. Have a look at [DOI: 10.5281/zenodo.7244918] and [DIO: 10.5281/zenodo.3583070] if you are interested in more details about the dataset you will be using. Please though note, that this dataset is only intended for training purposes in this course, so if you are planning to do project or thesis work on LCA, consult your supervisor to get access to suitable databases for your purpose._

## **LCA of Coal Power Production with and without Carbon Capture and Storage**

In this problem set, you will compare the environmental impacts from producing one kWh of coal based electricity with and without carbon capture and storage (CCS).

The inputs from the technosphere (background processes) to the coal power process required to produce 1 kWh of electricity without CCS are given in **Foreground_Background_list_PS11.csv**.

## **1. Flowcharts** _(0 points)_

Sketch two flowcharts of the inputs and outputs of the coal power process, one with and one without CCS. The emissions to air have to be represented as "Emissions to air" and the CCS has to be represented as well in the flowchart representing the inputs and outputs of the coal power process with CCS.

Even though this question will not be graded, it is strongly recommended to try to do a sketch on your own and ask questions during tutoring if there are any.

## **2. Environmental impacts of 1 kWh of electricity for both cases** _(60 points)_

Run the cell below to import `Aff`, `Abf`, `Afb` and `Abb` as Pandas DataFrames.

```python
Aff = convert_formulas(pd.read_csv(files_folder + 'Aff.csv', sep = ',', index_col = [0]))
Abf = convert_formulas(pd.read_csv(files_folder + 'Abf.csv', sep = ',', index_col = [0]))
Afb = convert_formulas(pd.read_csv(files_folder + 'Afb.csv', sep = ',', index_col = [0]))
Abb = convert_formulas(pd.read_csv(files_folder + 'Abb.csv', sep = ',', index_col = [0]))
```

Run the cell below to get the `C` matrix.

```python
C = pd.read_csv(files_folder + 'C.csv', sep = ',', index_col = [0])
```

Run the cell below to get the `Sb` matrix.

```python
Sb = pd.read_csv(files_folder + 'Sb.csv', sep = ',', index_col = [0])
```

#### **2a)** _(20 points)_ Get the `S` matrix

Define `Sf` and build the entire `S` matrix. In `S`, the stressors are in kg.

The direct stressors (in kg) to air per kWh of electricity are given in **Direct_stressors_list_PS11.csv** but the table is incomplete. Only the direct stressors in the case of electricity production by coal without CCS is given and you have to calculate the direct stressors for the electricity production with CCS using the following information.
- The coal power plant with CCS is assumed to have a CO2 capture efficiency of 90%.
- The extra CCS inputs lead to extra stressor outputs (due to the degradation of the solvent used to capture the CO2 in the flue gas). On a per kg CO2 captured basis, they are the following:
    - 35.5 mg NH3
    - 4.29 * 10^(-4) g NMVOC
- The CCS process also captures several other stressors, and on a per kg CO2 captured basis, they are the following:
    - 33.5 mg PM10
    - 36.6 mg NOx
    - 761 mg SOx 
    
For filling in **Direct_stressors_list_PS11.csv**, you have different options:
1. The easiest option for you to fill in the table is to copy-paste the table in an Excel document, do the calculations and copy-paste the results in **Direct_stressors_list_PS11.csv**. You need to be careful with the type of decimal separator used in your Excel (comma VS dot). The results in **Direct_stressors_list_PS11.csv** should be using dot as this is for Python.
2. You can also write formulas directly in **Direct_stressors_list_PS11.csv** but your formulas should not refer to other cells like "=C6" (as you would do in Excel), you need to write all the numbers you are using. It looks like it works in the csv file, but when imported to Python it does not work.
3. You can also use the pandas loc function to calculate the values of stressors.

```python
def fill_in_Sf_from_table(Sf, table):
    for i in Sf.index:
        for j in Sf.columns:
            if table[(table['Stressor'] == i) & (table['Foreground process'] == j)]['Value'].empty == False:
                Sf.loc[i, j] = table[(table['Stressor'] == i) & (table['Foreground process'] == j)]['Value'].values[0]
    return Sf

### This is importing the table
Sf_list = convert_formulas(pd.read_csv(files_folder + 'Direct_stressors_list_PS11.csv', sep = ',', index_col = [0, 1])).reset_index()
```

<!-- grade_id: cell-d6d71317c2a4426e, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e2bc5a40a2ae97d2, points: 20, grade: true, solution: false -->
```python
assert S.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (S.index == Sb.index).all(), 'The index do not correspond'
assert (S.columns == list(Aff.columns) + list(Abb.columns)).all(), 'The columns do not correspond'
```

#### **2b)** _(5 points)_ Fill in `Aff`

<!-- grade_id: cell-2562e8d25636c574, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-3130a835362333c2, points: 5, grade: true, solution: false -->
```python
assert Aff.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Aff.index == ['coal power without ccs', 'coal power with ccs']).all(), 'The index do not correspond'
assert (Aff.columns == ['coal power without ccs', 'coal power with ccs']).all(), 'The columns do not correspond'
```

#### **2c)** _(20 points)_ Fill in `Abf`

Use the file **Foreground_Background_list_PS11.csv** to fill in `Abf`. Inputs to the electricity production without CCS are given but the value of "Electrical machinery and apparatus n.e.c. (31) - tonnes" is missing. You have to calculate it knowing that the electricity production by coal requires 0.15 tonnes of "Electrical machinery and apparatus n.e.c. (31) - tonnes" per TJ of electricity produced.

You also have to add rows to the table to generate the inputs required for the electricity production with CCS using the following information:
- The CCS system requires additional inputs of coal: 15% more coal is required per kWh of electricity produced.
- There are extra operational inputs required to run the CCS system. They are the following (values are written on a per kg CO2 captured basis):
    - 1.60 g monoethanolamine as 'Chemicals nec - tonnes'
    - 0.13 g sodium hydroxide as 'Chemicals nec - tonnes'
    - 3.20 g reclaimer waste as 'Other waste for treatment: waste water treatment - tonnes'
    
For filling in **Foreground_Background_list_PS11.csv**, you have different options:
1. The easiest option for you to fill in the table is to copy-paste the table in an Excel document, do the calculations and copy-paste the results in **Foreground_Background_list_PS11.csv**. You need to be careful with the type of decimal separator used in your Excel (comma VS dot). The results in **Foreground_Background_list_PS11.csv** should be using dot as this is for Python.
2. You can also write formulas directly in **Foreground_Background_list_PS11.csv** but your formulas should not refer to other cells like "=C6" (as you would do in Excel), you need to write all the numbers you are using. It looks like it works in the csv file, but when imported to Python it does not work.
3. You can also use the pandas loc function to calculate the values of inputs.

```python
def fill_in_Abf_from_table(Abf, table):
    for i in Abf.index:
        for j in Abf.columns:
            if table[(table['Background process'] == i) & (table['Foreground process'] == j)]['Value'].empty == False:
                Abf.loc[i, j] = table[(table['Background process'] == i) & (table['Foreground process'] == j)]['Value'].values[0]
    return Abf

### This is importing the table
Abf_list = convert_formulas(pd.read_csv(files_folder + 'Foreground_Background_list_PS11.csv', sep = ',', index_col = [0, 1])).reset_index()
```

<!-- grade_id: cell-5e4e473b13be2fd2, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-7f47ce18ea29796f, points: 20, grade: true, solution: false -->
```python
assert Abf.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Abf.index == Abb.columns).all(), 'The index do not correspond'
assert (Abf.columns == Aff.columns).all(), 'The columns do not correspond'
```

#### **2d)** _(5 points)_ Fill in `Afb`, build the `A` matrix and calculate the `L` matrix

<!-- grade_id: cell-7c21777ed6e637a2, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-7ee44980afe3d251, points: 5, grade: true, solution: false -->
```python
assert (L.index == list(Aff.columns) + list(Abb.columns)).all(), 'The index do not correspond'
assert (L.columns == list(Aff.columns) + list(Abb.columns)).all(), 'The columns do not correspond'
```

#### **2e)** _(5 points)_ Calculate the output vectors `x_withoutCCS` and `x_withCCS`

Calculate `x_withoutCCS` and `x_withCCS` and rename their column into "Output". You have to define the two vectors for the final demand: (1) 1 kWh of electricity produced by coal without CCS and (2) 1 kWh of electricity produced by coal with CCS.

<!-- grade_id: cell-f385fd95f31b0e39, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-0b3ed46ec78b38e1, points: 5, grade: true, solution: false -->
```python
assert (x_withoutCCS.index == list(Aff.columns) + list(Abb.columns)).all(), 'The index do not correspond'
assert (x_withoutCCS.columns == ['Output']).all(), 'The columns do not correspond'
assert (x_withCCS.index == list(Aff.columns) + list(Abb.columns)).all(), 'The index do not correspond'
assert (x_withCCS.columns == ['Output']).all(), 'The columns do not correspond'
```

#### **2f)** _(5 points)_ Calculate `d_withoutCCS` and `d_withCCS` and rename their column into "Total impacts"

<!-- grade_id: cell-25c162b3444cf9d5, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-1b8c859c614dca7a, points: 5, grade: true, solution: false -->
```python
assert (d_withoutCCS.index == C.index).all(), 'The index do not correspond'
assert (d_withoutCCS.columns == ['Total impacts']).all(), 'The columns do not correspond'
assert (d_withCCS.index == C.index).all(), 'The index do not correspond'
assert (d_withCCS.columns == ['Total impacts']).all(), 'The columns do not correspond'
```

## **3. Comparison GWP100 for both cases** _(20 points)_

Compare the total GWP100 (Climate change) impacts for both cases. How much better is the CCS power plant?

1. The GHG emissions are reduced by 75%.
2. The GHG emissions are reduced by 57%.
3. The GHG emissions are reduced by 65%.

Define your answer as the variable `answer_3a` (e.g. `answer_3a = {1}` if you think the correct statement is statement {1} or `answer_3a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-8b557c22a345937f, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-9d68bd79f3d30a54, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_3a]), 'The answer(s) must be a set of values'
```

You have to fill in the table `GWP_comparison_direct_indirect` which is for breaking down the GWP impacts between direct (i.e. from operational coal power emissions) and indirect GWP impacts.

```python
GWP_comparison_direct_indirect = pd.DataFrame(index = ['direct impact', 'indirect impact'], columns = ['coal power without ccs', 'coal power with ccs'], dtype = float)
display(GWP_comparison_direct_indirect)
```

<!-- grade_id: cell-b61fba51fbba9688, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-20984f5ea9b91fa4, points: 10, grade: true, solution: false -->
```python
assert GWP_comparison_direct_indirect.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert np.round(GWP_comparison_direct_indirect.sum().loc['coal power without ccs']) == np.round(d_withoutCCS.loc["Climate change - GWP100 (kg CO2-eq)", 'Total impacts']), 'The sum of direct + indirect should be the same value as in d'
assert np.round(GWP_comparison_direct_indirect.sum().loc['coal power with ccs']) == np.round(d_withCCS.loc["Climate change - GWP100 (kg CO2-eq)", 'Total impacts']), 'The sum of direct + indirect should be the same value as in d'
assert (GWP_comparison_direct_indirect.index == ['direct impact', 'indirect impact']).all(), 'The index do not correspond'
assert (GWP_comparison_direct_indirect.columns == ['coal power without ccs', 'coal power with ccs']).all(), 'The columns do not correspond'
```

How much potential GWP mitigation is lost due to the indirect impacts of the extra inputs and waste disposal required? 
 
1. CCS reduces direct impacts by 90%.
2. CCS reduces direct impacts by 75%.
3. CCS reduces indirect impacts by 75%.
4. CCS increase indirect impacts by 19%.

Define your answer as the variable `answer_3b` (e.g. `answer_3b = {1}` if you think the correct statement is statement {1} or `answer_3b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-f1934291a1d95b7f, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-902bdc73ad1dc46a, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_3b]), 'The answer(s) must be a set of values'
```

## **4. Other impacts categories** _(20 points)_

Which one(s) of these statements are correct?

1. CCS is reducing impacts for all categories.
2. CCS is creating some problem shifting when it comes to the impacts generated by the electricity production.
3. Only for acidification, climate change and photochemical oxidation, the CCS case reduces the impacts compared to the case without CCS.
4. Human and freshwater toxicity increase by around 20% when the CCS is being used for the electricity production.

Define your answer as the variable `answer_4a` (e.g. `answer_4a = {1}` if you think the correct statement is statement {1} or `answer_4a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-d1d1720644f78c1b, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-898baff9cd3a6e1b, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_4a]), 'The answer(s) must be a set of values'
```

1. With CCS, NH3 contributing to human toxicity increases by around 680%.
2. With CCS, the CO2 contributing to climate change decreases by 90%.
3. Besides CO2, all the other stressors contributing to climate change increase when CCS is implemented.
4. NOx is the second largest contributor to acidification both for electricity produced without CCS and with CCS. 

Define your answer as the variable `answer_4b` (e.g. `answer_4b = {1}` if you think the correct statement is statement {1} or `answer_4b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-673f46ade6ba8e3b, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-1a3b53ae3e7c0fd6, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_4b]), 'The answer(s) must be a set of values'
```