# **Problem Set 12**

## **Tasks**
- Assembling A and S matrices for both cases
- Calculating d for both cases
- Answering the multiple choice questions


## **Passing requirement**
- **At least 80 points**

**_Please run the cell below before you start the assignment_**

```python
# Import required Python libraries
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns

# Folder for the files
files_folder = "files_PS12_V1_0/"

### In case formulas are entered in the csv files & correcting the datatypes
def convert_formulas(matrix):
    for row in matrix.index:
        for column in matrix.columns:
            if type(matrix.loc[row, column]) == str:
                if " " in matrix.loc[row, column]:
                    matrix.loc[row, column] = matrix.loc[row, column].lstrip(' ')
                if "=" in matrix.loc[row, column]:
                    matrix.loc[row, column] = matrix.loc[row, column].lstrip('=')
                matrix.loc[row, column] = pd.eval(matrix.loc[row, column])
    ### Get the data types right in the DataFrames
    s = matrix.select_dtypes(include='object').columns
    matrix[s] = matrix[s].astype("float64")
    return matrix
```

## **LCA of vehicle lightweighting**

In this problem set you will be analysing the environmental impact from lightweighting strategies in the automotive industry. A car’s glider (i.e. its body and chassis) is mostly made of steel, a relatively cheap but heavy material. It is technically possible to substitute a significant amount of this glider by another material, the main candidate being aluminium. However this material is more energy-intensive to produce than steel.

Our design is for a low-medium segment car, powered by a diesel engine. The car weighs **1450 kg** in its original version (before any lightweighting). Its body contains **660 kg** of steel ('Basic iron and steel and of ferro-alloys and first products thereof - tonnes'), of which **450 kg** can be substituted with another material. The car is assumed to have a lifetime of **240000 km**, and consumes **6.0 L/100 km**.

<img src="files_PS12_V1_0/car_body.jpg">

Two notions will be important in this Problem Set:
- The material substitution rate (MSR), a ratio that conveys how much new material can replace 1 kg of original material, e.g. aluminium has an MSR of 0.6 when compared with steel, meaning 600 g of aluminium can replace 1 kg of steel
- The fuel reduction value (FRV), a ratio that conveys how much fuel consumption is saved per 100 kg of lightweighting

|||
|-|-|
|**Strategy**|**MSR [1]**|
|Aluminium substitutes steel|0.6|

Ref [1]: Based on the UCSB Energy & GHG Model: https://www.worldautosteel.org/life-cycle-thinking/ucsb-energy-ghg-model/ 

The FRV for this type of vehicle is 0.3 L/100 km/100 kg. How to interpret this number: if the car weighed 1250 kg instead of 1450 kg (200 kg lighter), fuel savings would amount to 0.3 L/100 km * 200 kg/100 kg = 0.6 L/100 km, i.e. a new consumption of 5.4 L/100 km.

Diesel has an average density of 0.85 kg/L and its combustion emits 2.68 kg CO2/L. Therefore, the original car consumes 51 g of diesel per km, emitting 160.8 g CO2/km.

In Parts 2 and 3, the original car is being evaluated. The lightweighting strategy will be evaluated in Parts 4, 5 and 6.


_**Disclaimer**: We are using the hybrid version of EXIOBASE (a global, detailed Multi-Regional Environmentally Extended Supply-Use Table (MR-SUT) and Input-Output Table (MR-IOT)) to model background processes [DOI: 10.5281/zenodo.7244918] combined with stressors/impacts from EXIOBASE [DIO: 10.5281/zenodo.3583070]. For performing an LCA, it is recommended to use LCA databases such as Ecoinvent. These large databases are however often not open access. It was decided that for this course, we would stick to the matrix form of LCA using open data. If for further projects (e.g. semester project or master's thesis in the 2nd year of the master), a proper LCA has to be performed, you should contact one of the instructors from the course and we will provide you with the required tools._


## **1. Flowchart** _(0 points)_

Sketch a flowchart of the system with all the foreground components.

Even though this question will not be graded, it is strongly recommended to try to do a sketch on your own and ask questions during tutoring if there are any.

## **2. Data preparation** _(5 points)_

Run the cell below to import `Aff`, `Abf`, `Afb` and `Abb` as Pandas DataFrames.

```python
Aff = convert_formulas(pd.read_csv(files_folder + 'Aff.csv', sep = ',', index_col = [0]))
Abf = convert_formulas(pd.read_csv(files_folder + 'Abf.csv', sep = ',', index_col = [0]))
Afb = convert_formulas(pd.read_csv(files_folder + 'Afb.csv', sep = ',', index_col = [0]))
Abb = convert_formulas(pd.read_csv(files_folder + 'Abb.csv', sep = ',', index_col = [0]))
```

Run the cell below to get the `C` matrix.

```python
C = pd.read_csv(files_folder + 'C.csv', sep = ',', index_col = [0])
```

Run the cell below to get the `Sb` matrix.

```python
Sb = pd.read_csv(files_folder + 'Sb.csv', sep = ',', index_col = [0])
```

Define `Sf` (think if it should contain values or not) and build the entire `S` matrix. In `S`, the stressors are in kg.

<!-- grade_id: cell-07756fb685b586b6, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-029ed5a57437ea0b, points: 5, grade: true, solution: false -->
```python
assert S.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (S.index == Sb.index).all(), 'The index do not correspond'
assert (S.columns == list(Aff.columns) + list(Abb.columns)).all(), 'The columns do not correspond'
```

## **3. Environmental impacts of transportation over 1 km** _(30 points)_

#### **3a)** _(5 points)_ Fill in `Aff`

<!-- grade_id: cell-3f89b04c7a4a8570, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-6ea0430d27b788d0, points: 5, grade: true, solution: false -->
```python
assert Aff.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Aff.index == ['Transport, passenger car', 'Passenger car', 'Engine', 'Glider']).all(), 'The index do not correspond'
assert (Aff.columns == ['Transport, passenger car', 'Passenger car', 'Engine', 'Glider']).all(), 'The columns do not correspond'
```

#### **3b)** _(5 points)_ Fill in `Abf`

Use the file **Foreground_Background_list_PS12.csv** to fill in `Abf`. 

```python
def fill_in_Abf_from_table(Abf, table):
    for i in Abf.index:
        for j in Abf.columns:
            if table[(table['Background process'] == i) & (table['Foreground process'] == j)]['Value'].empty == False:
                Abf.loc[i, j] = table[(table['Background process'] == i) & (table['Foreground process'] == j)]['Value'].values[0]
    return Abf

### This is importing the table
Abf_list = pd.read_csv(files_folder + 'Foreground_Background_list_PS12.csv')
```

<!-- grade_id: cell-a547216f0fd4cf83, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-d9029e98dffb28b2, points: 5, grade: true, solution: false -->
```python
assert Abf.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Abf.index == Abb.columns).all(), 'The index do not correspond'
assert (Abf.columns == Aff.columns).all(), 'The columns do not correspond'
```

#### **3c)** _(5 points)_ Fill in `Afb`, build the `A` matrix and calculate the `L` matrix

<!-- grade_id: cell-b1d21839fad23797, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-3d61090acbd1a1e8, points: 5, grade: true, solution: false -->
```python
assert (L.index == A.columns).all(), 'The index do not correspond'
assert (L.columns == A.columns).all(), 'The columns do not correspond'
```

#### **3d)** _(5 points)_ Calculate the output vector `x`

Calculate `x` and rename its column into "Output". You have to define the vector for the final demand: 1 km of transportation by passenger car.

<!-- grade_id: cell-910f57451f2b737f, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-021a315ed3b25e11, points: 5, grade: true, solution: false -->
```python
assert (x.index == A.index).all(), 'The index do not correspond'
assert (x.columns == ['Output']).all(), 'The columns do not correspond'
```

#### **3e)** _(5 points)_ Calculate `d` and rename its column into "Total impacts"

<!-- grade_id: cell-9c67482254412021, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e2ce50d7838508e4, points: 5, grade: true, solution: false -->
```python
assert (d.index == C.index).all(), 'The index do not correspond'
assert (d.columns == ['Total impacts']).all(), 'The columns do not correspond'
```

Which statement is correct?
 
1. The direct greenhouse gas emissions represent about 51 % of the lifecycle greenhouse gas emissions of 1 km of "Transport, passenger car".
2. The direct greenhouse gas emissions represent about 85 % of the lifecycle greenhouse gas emissions of 1 km of "Transport, passenger car".
3. The direct greenhouse gas emissions represent about 69 % of the lifecycle greenhouse gas emissions of 1 km of "Transport, passenger car".

Define your answer as the variable `answer_3e` (e.g. `answer_3e = {1}` if you think the correct statement is statement {1}).

<!-- grade_id: cell-ebabeda2feb4a20f, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-71024178f0a3fd0b, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_3e]), 'The answer(s) must be a set of values'
```

## **4. Lightweighting strategies** _(25 points)_

Now let us model the lightweight car with the strategy of aluminium substitution. Given the information in the introduction, fill in the table `df_strategy` if you replace 450 kg of steel with aluminium.

You have to calculate:
- how much new material is needed (Substitution material (in kg)), 
- how much weight is saved - it should be a value > 0 (Weight saving (in kg)), 
- the new consumption of the car (New fuel consumption (in kg/km)),
- the new emissions (New CO2 emissions (in kgCO2/km)). 

```python
### This is the DataFrame you have to fill in
df_strategy = pd.DataFrame(columns = ['Substitution material (in kg)', 
                                      'Weight saving (in kg)', 
                                      'New fuel consumption (in kg/km)', 
                                      'New CO2 emissions (in kgCO2/km)'], 
                           index = ['Aluminium subsitution'],
                          dtype = float)

display(df_strategy)
```

<!-- grade_id: cell-c43207972f88159c, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-5ad4c1f298e09012, points: 20, grade: true, solution: false -->
```python
assert df_strategy.isnull().values.any() == False, 'There are still NaN values in the matrix'
```

How much does the car weight with the aluminium strategy? Define your answer as `new_weight_alu`.

<!-- grade_id: cell-68976282b7e39011, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-1063dc276bc61c03, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [int] for i in [new_weight_alu]), 'The answer(s) must be an integer'
```

## **5. Lightweighting - impacts calculations** _(30 points)_

We want to evaluate the impacts of substituting 450 kg of steel with aluminium.

#### **5a)** _(20 points)_ Fill in `Abf_alu`

Only part of the steel ('Basic iron and steel and of ferro-alloys and first products thereof - tonnes') from the glider has be substituted with aluminium ('Aluminium and aluminium products - tonnes'). Do not forget to change all the inputs which are required to change.

```python
### Abf_alu is a copy of Abf and this is the DataFrame you have to modify. Otherwise, if you use Abf directly, you will change its values.
Abf_alu = Abf.copy()
```

<!-- grade_id: cell-4c270abc1bb4e331, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-349940f649a8c105, points: 20, grade: true, solution: false -->
```python
assert (Abf_alu.index == Abb.columns).all(), 'The index do not correspond'
assert (Abf_alu.columns == Aff.columns).all(), 'The columns do not correspond'
```

#### **5b)** _(10 points)_ Calculate `d_alu` and rename its column into "Total impacts"

Do not forget to change all the variables of the system which require to be changed to calculate the impacts of the passenger car after substitution with aluminium!

<!-- grade_id: cell-309c0ab535313706, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-1aa4b9bbfacf367e, points: 10, grade: true, solution: false -->
```python
assert (d_alu.index == C.index).all(), 'The index do not correspond'
assert (d_alu.columns == ['Total impacts']).all(), 'The columns do not correspond'
```

## **6. Comparison GWP between original car and car with aluminium substitution** _(10 points)_

By how much % has changed the GWP between the original car and car with aluminium substitution? 

1. It is increasing by 5%.
2. It is increasing by 3%.
3. It is decreasing by 1%.
4. It is decreasing by 4%. 

Define your answer as the variable `answer_6a` (e.g. `answer_6a = {1}` if you think the correct statement is statement {1}).

<!-- grade_id: cell-e99f61e176f169f3, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-9d69dabe6bfb91e8, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6a]), 'The answer(s) must be a set of values'
```

Given the results you obtained, which statement is correct?
 
1. The direct greenhouse gas emissions represent about 59 % of the lifecycle greenhouse gas emissions of 1 km of "Transport, passenger car".
2. The direct greenhouse gas emissions represent about 63 % of the lifecycle greenhouse gas emissions of 1 km of "Transport, passenger car".
3. The direct greenhouse gas emissions represent about 73 % of the lifecycle greenhouse gas emissions of 1 km of "Transport, passenger car".

Define your answer as the variable `answer_6b` (e.g. `answer_6b = {1}` if you think the correct statement is statement {1}).

<!-- grade_id: cell-0568c2cd19d4abac, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-abb47cd2df542694, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6b]), 'The answer(s) must be a set of values'
```

Given the following greenhouse gas emissions for aluminium and steel, and your answer to the previous question, what do you think about the change in greenhouse gas emissions? 

_This is an open question for you to reflect on the results you are getting. Your answer is not graded, but it is highly recommended to ask questions during the tutoring session if anything is unclear._

<!-- grade_id: cell-15f1ba06acc19d4a, points: 0, grade: false, solution: true -->
```python
### GWP per tonne of aluminium
y_aluminium = pd.DataFrame(index = list(Aff.index) + list(Abf.index), columns = ["Final demand"], dtype = float).fillna(0)
y_aluminium.loc['Aluminium and aluminium products - tonnes'] = 1
x_aluminium = L @ y_aluminium
d_aluminium = C @ S @ x_aluminium
d_aluminium = d_aluminium.rename(columns={'Final demand': 'Total impacts'})
display(d_aluminium.loc['Climate change - GWP100 (kg CO2-eq)'])

### GWP per tonne of steel
y_steel = pd.DataFrame(index = list(Aff.index) + list(Abf.index), columns = ["Final demand"], dtype = float).fillna(0)
y_steel.loc['Basic iron and steel and of ferro-alloys and first products thereof - tonnes'] = 1
x_steel = L @ y_steel
d_steel = C @ S @ x_steel
d_steel = d_steel.rename(columns={'Final demand': 'Total impacts'})
display(d_steel.loc['Climate change - GWP100 (kg CO2-eq)'])

# Write your answer here.
```