# **Problem Set 1 - Problem A**

## **Passing requirements**
**All of the following must be satisfied for a passing grade**
- Correctly filling in the process diagrams
- Correctly set up the A and L matrices and correctly answer the multiple choice questions for interpretation

**_Please run the cell below before you start the assignment_**

```python
# Import required Python libraries
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
```

This problem set is an introduction to the use of Python for the LCA assignements. We are using Pandas DataFrames as a way to code matrices and perform calculations. Pandas DataFrames are 2 dimensional data structure with labeled rows (index) and columns. The full Pandas documentation can be found here: https://pandas.pydata.org/docs/

The Python functions you might need to use in the assigments throughout the semester are the following:
- Read the indices of a DataFrame `df`: `df.index`
- Read the columns of a DataFrame `df`: `df.columns`
- Fill in a value from an index (row) and column of a DataFrame `df`: `df.loc[index, column] = value`
- Replace NaN values by 0 in a DataFrame `df`: `df = df.fillna(0)`
- Display a DataFrame `df`: `display(df)`
- Concatenate two DataFrames `df1` and `d2`:
    - Along the index: `pd.concat([df1, df2], axis = 0)`
    - Along the columns: `pd.concat([df1, df2], axis = 1)`
- Define an identity matrix `I` of size s: `I = np.identity(s)`
- Calculate the determinant of a matrix `M`: `np.linalg.det(M)`
- Inverse a DataFrame `df`: `np.linalg.inv(df)` -> /!\ this returns a `np.array` (numpy array - more here: https://numpy.org/doc/stable/reference/generated/numpy.array.html)
- Convert a `np.array` to a DataFrame with specified index and columns: `pd.DataFrame(np.array, index, columns)`
- Multiply two DataFrames `df1` and `df2`: `df3 = df1 @ df2`
    - `df1` and `df2` should have matching `df1.columns` with `df2.index` for the multiplication to run
    - After the multiplication, the index of `df3` are the index of `df1` and the columns of `df3` are the columns of `df2`
    - It is therefore necessary to rename index/columns of `df3` to get the correct names
- Rename index of a DataFrame `df`: `df = df.rename(index = {initial_index_name : new_index_name})`
- Rename column of a DataFrame `df`: `df = df.rename(columns = {initial_column_name : new_column_name})`

## **Leontief model**

For this problem, we look at only two production processes to illustrate the concept of inter-process connectivity, which are labeled Process A and B.

### **Production process i)** _(40 points)_

**Process A requires 0.04 units for its own production process and 0.6 units from Process B, and similarly, Process B requires 0.7 units from Process A and 0.02 units from itself.**

#### **1i) Simple process diagram** _(5 points)_

The diagram below represents the various inputs and outputs between the two nodes, demonstrating the inter-connectivity. Your task is to assign the inputs and outputs values.

<img src="files_PS1A_V1_0/ProblemA_diagram.PNG" width="700">

Replace **a**, **b**, **c** and **d** by the inputs and outputs values in the description of Production process i). Define your answers as the variables `answer_1i_a`, `answer_1i_b`, `answer_1i_c` and `answer_1i_d` (e.g. `answer_1i_a = 1` if you think the correct value to replace **a** is 1).

<!-- grade_id: cell-892862e95f72609a, points: 0, grade: false, solution: true -->
```python
answer_1i_a=0.04
answer_1i_b=0.7
answer_1i_c=0.02
answer_1i_d=0.6
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-222b4610a816bbfe, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [float, float, float, float] for i in [answer_1i_a, answer_1i_b, answer_1i_c, answer_1i_d]), 'The answer(s) must be a float'
```

#### **2i) A matrix** _(10 points)_

Your task is to set up the A matrix that is named `A_i`: the DataFrame is created for you, you have to fill it in.

```python
A_i = pd.DataFrame(index = ['A', 'B'], columns = ['A', 'B'], dtype = 'float64')
```

Fill in `A_i`

<!-- grade_id: cell-d3612f000304e1fd, points: 0, grade: false, solution: true -->
```python
A_i.loc['A', 'A'] = answer_1i_a
A_i.loc['B', 'A'] = answer_1i_d
A_i.loc['A', 'B'] = answer_1i_b
A_i.loc['B', 'B'] = answer_1i_c
display(A_i)
```

Output:
```
A     B
A  0.04  0.70
B  0.60  0.02
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-a1ba6f373a4b1dc3, points: 10, grade: true, solution: false -->
```python
assert A_i.isnull().values.any() == False, 'There are still NaN values in the matrix'
```

#### **3i) I-A** _(5 points)_

Define the identity matrix `I_i` and calculate the determinant `detIA_i` of the matrix (I-A) in the cell under.

<!-- grade_id: cell-0efb9e1dd1442d01, points: 0, grade: false, solution: true -->
```python
I_i=np.identity(2)
detIA_i=np.linalg.det(I_i-A_i)
print(detIA_i)
```

Output:
```
0.5208
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-64e0a31eefeccb8c, points: 3, grade: true, solution: false -->
```python
assert type(detIA_i) == np.float64, 'The answer(s) must be numeric'
```

State whether or not the matrix `A_i` is invertible. If the matrix `A_i` is invertible, write `invertible_i = True`. If it is not invertible, write `invertible_i = False`.

<!-- grade_id: cell-c16a2802c084a020, points: 0, grade: false, solution: true -->
```python
invertible_i = True
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-931051f7b57a3290, points: 2, grade: true, solution: false -->
```python
assert type(invertible_i) == bool, 'The answer(s) must be True/False'
```

#### **4i) Leontief Inverse matrix L** _(10 points)_

Determine the L matrix that you will name `L_i`. Remember that `L_i` must be a Pandas DataFrame.

<!-- grade_id: cell-be9fdd9f203211ac, points: 0, grade: false, solution: true -->
```python
L_i_arr=np.linalg.inv(I_i-A_i)
L_i=pd.DataFrame(L_i_arr, index = ['A', 'B'], columns = ['A', 'B'])
print(L_i)
```

Output:
```
A         B
A  1.881720  1.344086
B  1.152074  1.843318
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e2dc8455244e8a8b, points: 10, grade: true, solution: false -->
```python
assert (L_i.index == ['A', 'B']).all(), 'The index do not correspond'
assert (L_i.columns == ['A', 'B']).all(), 'The columns do not correspond'
```

#### **5i) Output vector x** _(10 points)_

Calculate the output vector x for four cases of the functional unit y matrix.

```python
y_i_1 = pd.DataFrame([1, 0], columns = ['Final demand'], index = ['A', 'B'])
y_i_2 = pd.DataFrame([0, 1], columns = ['Final demand'], index = ['A', 'B'])
y_i_3 = pd.DataFrame([1, 1], columns = ['Final demand'], index = ['A', 'B'])
y_i_4 = pd.DataFrame([3, 1], columns = ['Final demand'], index = ['A', 'B'])
```

Calculate `x_i_1` using `y_i_1`. Rename the column of `x_i_1` as "Output".

<!-- grade_id: cell-b8dc3797a568b242, points: 0, grade: false, solution: true -->
```python
x_i_1=L_i @ y_i_1
x_i_1 = x_i_1.rename(columns = {'Final demand' : 'Output'})
print(x_i_1)
display(x_i_1)
```

Output:
```
Output
A  1.881720
B  1.152074
     Output
A  1.881720
B  1.152074
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-7c688f0d2501e9e4, points: 2.5, grade: true, solution: false -->
```python
assert (x_i_1.columns == ['Output']).all(), 'The columns do not correspond'
```

Calculate `x_i_2` using `y_i_2`. Rename the column of `x_i_2` as "Output".

<!-- grade_id: cell-7cf3f9a181820c88, points: 0, grade: false, solution: true -->
```python
x_i_2=L_i @ y_i_2
x_i_2 = x_i_2.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-57d1a4f0dc23ad26, points: 2.5, grade: true, solution: false -->
```python
assert (x_i_2.columns == ['Output']).all(), 'The columns do not correspond'
```

Calculate `x_i_3` using `y_i_3`. Rename the column of `x_i_3` as "Output".

<!-- grade_id: cell-3d5acbefa3233f84, points: 0, grade: false, solution: true -->
```python
x_i_3=L_i @ y_i_3
x_i_3 = x_i_3.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-210444b95e4fb4b5, points: 2.5, grade: true, solution: false -->
```python
assert (x_i_3.columns == ['Output']).all(), 'The columns do not correspond'
```

Calculate `x_i_4` using `y_i_4`. Rename the column of `x_i_4` as "Output".

<!-- grade_id: cell-f85686ebc589f3df, points: 0, grade: false, solution: true -->
```python
x_i_4=L_i @ y_i_4
x_i_4 = x_i_4.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-9fc77f25911e147e, points: 2.5, grade: true, solution: false -->
```python
assert (x_i_4.columns == ['Output']).all(), 'The columns do not correspond'
```

### **Production process ii)** _(40 points)_

**Process A requires 0.15 units for its own production process and 0.4 unit from Process B, and similarly, Process B requires 0.35 from Process A and 0.09 from itself.**

#### **1ii) Simple process diagram** _(5 points)_

The diagram below represents the various inputs and outputs between the two nodes, demonstrating the inter-connectivity. Your task is to assign the inputs and outputs values.

<img src="files_PS1A_V1_0/ProblemA_diagram.PNG" width="700">

Replace **a**, **b**, **c** and **d** by the inputs and outputs values in the description of Production process ii). Define your answers as the variables `answer_1ii_a`, `answer_1ii_b`, `answer_1ii_c` and `answer_1ii_d` (e.g. `answer_1ii_a = 1` if you think the correct value to replace **a** is 1).

<!-- grade_id: cell-e86e49d5693381c7, points: 0, grade: false, solution: true -->
```python
answer_1ii_a = 0.15
answer_1ii_b = 0.35
answer_1ii_c = 0.09
answer_1ii_d = 0.4
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-92b4325ec60d7fba, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [float, float, float, float] for i in [answer_1ii_a, answer_1ii_b, answer_1ii_c, answer_1ii_d]), 'The answer(s) must be a float'
```

#### **2ii) A matrix** _(10 points)_

Your task is to set up the A matrix that is named `A_ii`: the DataFrame is created for you, you have to fill it in.

```python
A_ii = pd.DataFrame(index = ['A', 'B'], columns = ['A', 'B'], dtype = 'float64')
```

Fill in `A_ii`

<!-- grade_id: cell-d04ba704f39a57d2, points: 0, grade: false, solution: true -->
```python
A_ii.loc['A', 'A'] = answer_1ii_a
A_ii.loc['B', 'A'] = answer_1ii_d
A_ii.loc['A', 'B'] = answer_1ii_b
A_ii.loc['B', 'B'] = answer_1ii_c
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-866a3b576554772a, points: 10, grade: true, solution: false -->
```python
assert A_ii.isnull().values.any() == False, 'There are still NaN values in the matrix'
```

#### **3ii) I-A** _(5 points)_

Define the identity matrix `I_ii` and calculate the determinant `detIA_ii` of the matrix (I-A) in the cell under.

<!-- grade_id: cell-579aa3a13886f786, points: 0, grade: false, solution: true -->
```python
I_ii=np.identity(2)
detIA_ii=np.linalg.det(I_ii-A_ii)
print(detIA_ii)
```

Output:
```
0.6335000000000001
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-5f436974f98d79f3, points: 3, grade: true, solution: false -->
```python
assert type(detIA_ii) == np.float64, 'The answer(s) must be numeric'
```

State whether or not the matrix `A_ii` is invertible. If the matrix `A_ii` is invertible, write `invertible_ii = True`. If it is not invertible, write `invertible_ii = False`.

<!-- grade_id: cell-5440d9da3724e373, points: 0, grade: false, solution: true -->
```python
invertible_ii = True
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-b85d253bac8928f6, points: 2, grade: true, solution: false -->
```python
assert type(invertible_ii) == bool, 'The answer(s) must be True/False'
```

#### **4ii) Leontief Inverse matrix L** _(10 points)_

Determine the L matrix that you will name `L_ii`. Remember that `L_ii` must be a Pandas DataFrame.

<!-- grade_id: cell-81995f155fcff389, points: 0, grade: false, solution: true -->
```python
L_ii_arr=np.linalg.inv(I_ii-A_ii)
L_ii=pd.DataFrame(L_ii_arr, index = ['A', 'B'], columns = ['A', 'B'])
display(L_ii)
```

Output:
```
A         B
A  1.436464  0.552486
B  0.631413  1.341752
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-ce978bdd7149be1e, points: 10, grade: true, solution: false -->
```python
assert (L_ii.index == ['A', 'B']).all(), 'The index do not correspond'
assert (L_ii.columns == ['A', 'B']).all(), 'The columns do not correspond'
```

#### **5ii) Output vector x** _(10 points)_

Calculate the output vector x for four cases of the functional unit y matrix.

```python
y_ii_1 = pd.DataFrame([1, 0], columns = ['Final demand'], index = ['A', 'B'])
y_ii_2 = pd.DataFrame([0, 1], columns = ['Final demand'], index = ['A', 'B'])
y_ii_3 = pd.DataFrame([1, 1], columns = ['Final demand'], index = ['A', 'B'])
y_ii_4 = pd.DataFrame([3, 1], columns = ['Final demand'], index = ['A', 'B'])
```

Calculate `x_ii_1` using `y_ii_1`. Rename the column of `x_ii_1` as "Output".

<!-- grade_id: cell-296e543889ead99a, points: 0, grade: false, solution: true -->
```python
x_ii_1=L_ii @ y_ii_1
x_ii_1 = x_ii_1.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-9067afca8d9f163a, points: 2.5, grade: true, solution: false -->
```python
assert (x_ii_1.columns == ['Output']).all(), 'The columns do not correspond'
```

Calculate `x_ii_2` using `y_ii_2`. Rename the column of `x_ii_2` as "Output".

<!-- grade_id: cell-04ad58efdb0ec2be, points: 0, grade: false, solution: true -->
```python
x_ii_2=L_ii @ y_ii_2
x_ii_2 = x_ii_2.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-9688282055172993, points: 2.5, grade: true, solution: false -->
```python
assert (x_ii_2.columns == ['Output']).all(), 'The columns do not correspond'
```

Calculate `x_ii_3` using `y_ii_3`. Rename the column of `x_ii_3` as "Output".

<!-- grade_id: cell-9526697f05ebd929, points: 0, grade: false, solution: true -->
```python
x_ii_3=L_ii @ y_ii_3
x_ii_3 = x_ii_3.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-ac8495f71bd874bf, points: 2.5, grade: true, solution: false -->
```python
assert (x_ii_3.columns == ['Output']).all(), 'The columns do not correspond'
```

Calculate `x_ii_4` usingr `y_ii_4`. Rename the column of `x_ii_4` as "Output". 

<!-- grade_id: cell-3102cf4928eded73, points: 0, grade: false, solution: true -->
```python
x_ii_4=L_ii @ y_ii_4
x_ii_4 = x_ii_4.rename(columns = {'Final demand' : 'Output'})
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-ab317f84507ab0a7, points: 2.5, grade: true, solution: false -->
```python
assert (x_ii_4.columns == ['Output']).all(), 'The columns do not correspond'
```

### **Interpretation questions** _(20 points)_

#### **6a)** _(5 points)_ Which one(s) of these statement(s) is/are correct?

1. The first column of A is the "cooking recipe" of product A.
2. The second column of A is the "cooking recipe" of product B.
3. The first row of A is the "cooking recipe" of product A.
4. The second row of A is the "cooking recipe" of product B.
 
Define your answer as the variable `answer_6a` (e.g. `answer_6a = {1}` if you think the correct statement is statement {1} and `answer_6a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-0994745c9af8acb6, points: 0, grade: false, solution: true -->
```python
answer_6a = {1}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-fd8c79fd13a14458, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6a]), 'The answer(s) must be a set of values'
```

#### **6b)** _(5 points)_ Which one(s) of these statement(s) is/are correct?

For a matrix A with coefficients aij (i = row, j = column):
1. aii = requirement from process i to produce 1 unit of process i
2. aij = requirement from process j to produce 1 unit of process i
3. aij = requirement from process i to produce 1 unit of process j

Define your answer as the variable `answer_6b` (e.g. `answer_6b = {1}` if you think the correct statement is statement {1} and `answer_6b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-daa96a6a959f868b, points: 0, grade: false, solution: true -->
```python
answer_6b = {1, 3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-bf55529e43f5f19f, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6b]), 'The answer(s) must be a set of values'
```

#### **6c)** _(5 points)_ Which one(s) of these statement(s) is/are correct?

det(I-A):
1. Tells us if the matrix (I-A) is invertible when it is != 0
2. Tells us if the system contains self-sustaining processes when it is > 0
3. Tells us if the system satisfies the Hawkins-Simon condition when it is > 0

Define your answer as the variable `answer_6c` (e.g. `answer_6c = {1}` if you think the correct statement is statement {1} and `answer_6c = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-fa740b8874a45df5, points: 0, grade: false, solution: true -->
```python
answer_6c = {1}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-05c4582c5e998bf8, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6c]), 'The answer(s) must be a set of values'
```

#### **6d)** _(5 points)_ Which one(s) of these statement(s) is/are correct?

1. The columns of L represent the total output per unit final demand of a given process. 
2. When there is only one unit of final demand of a given process, we see that only the row of L that represents the process requirements for the process shows up in the x vector.
3. lij is the total output from process j that is required to produce a unit final demand of process i.
4. lij is the total output from process i that is required to produce a unit final demand of process j.

Define your answer as the variable `answer_6d` (e.g. `answer_6d = {1}` if you think the correct statement is statement {1} and `answer_6d = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-6bc2dc8e57e242c2, points: 0, grade: false, solution: true -->
```python
answer_6d = {1, 4}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-5dad984158ba0da9, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6d]), 'The answer(s) must be a set of values'
```