# **Problem Set 2**

**This problem set is divided into two parts: Part A and Part B. For getting this problem set approved, you need at least 160 points (out of 200 points).**

# **Problem Set 2 - Part a**

## **Passing requirements**

**All of the following must be satisfied for a passing grade**
- Correct understanding of the flowchart
- Correctly assembled A and S matrices
- Correctly calculated L matrix and x vector
- Correct interpretation of the matrices

**_Please run the cell below before you start the assignment_**

```python
# Import required Python libraries
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns

# Folder for the files
files_folder = "files_PS2_V1_0/"
```

## **Internal Combustion Engine vehicle**

In this exercise, we will proceed with the ICE vehicle construction example and practice modeling the structure of the requirements matrix for the whole system, including both the foreground and background system.

You are given templates for the construction of the foreground system, a 10 x 10 background data set collected from an LCA database (`Abb`, `Sb`) and a characterization matrix (`C`) found in the folder **files_PS2_V1_0**. 

**Description of the ICE vehicle system**

The ICE vehicle consists of four parts: the body, the fuel tank, the powertrain, and the gearbox.

- The **vehicle’s body** component requires inputs of 300 kg steel, 70 kg aluminum, 65 kg plastic, and 8110 kWh electricity to assemble. Facility heating requirements are met using the same heating system as for the vehicle assembly, and each body manufactured requires 454 MJ heat from LFO and 3110 MJ heat from natural gas. Emissions of 88 grams CH$_4$, 1.74 tons CO$_2$, 38.6 kg NO$_x$, 347 grams PM$_{10}$, and 18.1 kg SO$_2$ are generated. 
- The **fuel tank** requires 3.2 kg plastic and 50.5 kWh electricity to manufacture. Tank manufacturing generates direct emissions of 7.1 grams CH$_4$, 135 kg CO$_2$, 612 grams NO$_x$, 3.1 grams PM$_{10}$, and 92 grams SO$_2$. 
- The **powertrain** requires inputs of 40 kg steel, 32 kg aluminum, 2.3 kg plastic, and 2240 kWh electricity. Each powertrain manufactured requires 518 MJ heat from LFO and 2456 MJ heat from a natural gas furnace. Emissions of 120 grams CH$_4$, 1.48 tons CO$_2$, 21.1 kg NO$_x$, 412 grams PM$_{10}$, and 14.4 kg SO$_2$ are generated. 
- Lastly, the **vehicle’s gearbox** requires 24 kg steel, 4.6 kg aluminum, and 610 kWh electricity to assemble. Producing the gearbox generates emissions of 28 grams CH$_4$, 889 kg CO$_2$, 4.3 kg NO$_x$, 13 grams PM$_{10}$ and 11 grams SO$_2$.

The manufacturing plant for the ICE vehicle requires electricity inputs of 100 000 kWh electricity for the complete assembly. At the manufacturing plant, the vehicle’s windows, consisting of 42 kg glass, are installed. Facility heating is also needed at the plant. This heating is produced by a 1 MW industrial furnace combusting light fuel oil (LFO) as well as from a 100 kW natural gas furnace. For each vehicle assembled at the plant, 2320 MJ heat from LFO and 9061 MJ heat from natural gas are needed. The combustion of the heating fuels during the vehicle assembly process generates direct emissions of 5100 grams CH$_4$, 9.42 tons CO$_2$, 156.3 kg NO$_x$, 9.81 kg PM$_{10}$ and 38.7 kg SO$_2$. 

In order to drive the new vehicle out from the final manufacturing facility to be loaded onto freight trucks, each vehicle is given 2 liters of diesel. Each liter of diesel has an energy density of approximately 37 MJ. Diesel and LFO can be considered to be the same product.

## **1. Flowchart of the ICE vehicle system** _(20 points)_

The flowchart (shown below) is representing the production of one ICE vehicle, including both the foreground and the background systems. The flowchart is however incomplete and your task is to find the missing components by answering the cells under the figure. The flowchart is on purpose not showing the background flows.

![PS2_flowchart](files_PS2_V1_0/PS2_flowchart.PNG)

#### **1a)** _(5 points)_ Replace A, B and C by either Aff, Abf, Afb or Abb.
Define your answer as the variables `answer_1a_A`, `answer_1a_B` and `answer_1a_C` (e.g. `answer_1a_A = "Aff"` with `"Aff"` as a string if you think the correct answer is Aff).

In Aff, Abf, Afb or Abb:
- ff means “foreground to foreground” 
- fb means “foreground to background” 
- bf means “background to foreground” 
- bb means "background to background” 

<!-- grade_id: cell-295d59142a1b4b70, points: 0, grade: false, solution: true -->
```python
answer_1a_A = "Abb"
answer_1a_B = "Abf"
answer_1a_C = "Aff"
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-d7fc4e2a4933faaa, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [str, str, str] for i in [answer_1a_A, answer_1a_B, answer_1a_C]), 'The answer(s) must be a string'
```

#### **1b)** _(10 points)_ Replace D, E, F, G by their values and units

Define your answer as the variables `answer_1b_D`, `answer_1b_E`, `answer_1b_F` and `answer_1b_G` (e.g. `answer_1b_D = "1kg"` with `"1kg"` as a string and without space if you think the correct answer is 1 kg).

<!-- grade_id: cell-72cf4dea1c2580ab, points: 0, grade: false, solution: true -->
```python
answer_1b_D = "9061MJ"
answer_1b_E = "2394MJ"
answer_1b_F= "100000kWh"
answer_1b_G = "42kg"
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-4a35e37b14a82140, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [str, str, str, str] for i in [answer_1b_D, answer_1b_E, answer_1b_F, answer_1b_G]), 'The answer(s) must be a string'
```

#### **1c)** _(5 points)_ Replace H by its value

Define your answer as the variable `answer_1c_H` (e.g. `answer_1b_H = 0` if you think the correct answer is 0).

<!-- grade_id: cell-10b457955f552d97, points: 0, grade: false, solution: true -->
```python
answer_1c_H = 1
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-294e904bcd8b89b0, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [int] for i in [answer_1c_H]), 'The answer(s) must be an integer'
```

## **2. A and S matrices** _(40 points)_

Your task is to set up the complete `A` matrix and stressor matrix `S` using values from the description of the ICE vehicle. 
`Abb` and `Sb` are already given. You have to fill in the matrices `Aff`, `Abf`, `Afb` and `Sf`. For that, you have two options: 
- Option 1: filling in the matrices in the csv files located in the **files_PS2_V1_0** folder and save the csv files by using "ctrl+s"
- Option 2: filling in the matrices in the cells in the sections 2a), 2b), 2c) and 2e). If you go for that option, we recommend you to make a quick hand sketch of the matrices to visualize how they should look like.

**Do not hesitate to use the display() function to check you have built the DataFrames as you expect them to be!**

### **A matrix** _(20 points)_

Run the cell below to import `Aff`, `Abf`, `Afb` and `Abb` as Pandas DataFrames. Do not change the names of the index or columns of the matrices nor their order as the autograding process might return an incorrect answer.

```python
Aff = pd.read_csv(files_folder + 'Aff.csv', sep = ',', index_col = [0])
Abf = pd.read_csv(files_folder + 'Abf.csv', sep = ',', index_col = [0])
Afb = pd.read_csv(files_folder + 'Afb.csv', sep = ',', index_col = [0])
Abb = pd.read_csv(files_folder + 'Abb.csv', sep = ',', index_col = [0])
```

#### **2a)** _(5 points)_ Fill in `Aff`

<!-- grade_id: cell-346e751e9390be82, points: 0, grade: false, solution: true -->
```python
Aff.loc['Body (piece)', 'Car (piece)'] = 1
Aff.loc['Tank (piece)', 'Car (piece)'] = 1
Aff.loc['Powertrain (piece)', 'Car (piece)'] = 1
Aff.loc['Gearbox (piece)', 'Car (piece)'] = 1
Aff = Aff.fillna(0)
display(Aff)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Car (piece)                 0.0           0.0           0.0   
Body (piece)                1.0           0.0           0.0   
Tank (piece)                1.0           0.0           0.0   
Powertrain (piece)          1.0           0.0           0.0   
Gearbox (piece)             1.0           0.0           0.0   

                    Powertrain (piece)  Gearbox (piece)  
Car (piece)                        0.0              0.0  
Body (piece)                       0.0              0.0  
Tank (piece)                       0.0              0.0  
Powertrain (piece)                 0.0              0.0  
Gearbox (piece)                    0.0              0.0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-cee0162ae5a7cdf5, points: 5, grade: true, solution: false -->
```python
assert Aff.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Aff.index == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                      'Gearbox (piece)']).all(), 'The index do not correspond'
assert (Aff.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)']).all(), 'The columns do not correspond'
```

#### **2b)** _(5 points)_ Fill in `Abf`

<!-- grade_id: cell-c097d5208f9cbaeb, points: 0, grade: false, solution: true -->
```python
Abf.loc['Electricity (N. Gas kWh)', 'Car (piece)'] = 100000
Abf.loc['N. Gas (MJ)', 'Car (piece)'] = 9061
Abf.loc['Glass (kg)', 'Car (piece)'] = 42
Abf.loc['Diesel (MJ)', 'Car (piece)'] = 2394

Abf.loc['Diesel (MJ)', 'Body (piece)'] = 454
Abf.loc['Electricity (N. Gas kWh)', 'Body (piece)'] = 8110
Abf.loc['Steel (kg)', 'Body (piece)'] = 300
Abf.loc['Aluminium (kg)', 'Body (piece)'] = 70
Abf.loc['Plastic (kg)', 'Body (piece)'] = 65
Abf.loc['N. Gas (MJ)', 'Body (piece)'] = 3110

Abf.loc['Electricity (N. Gas kWh)', 'Tank (piece)'] = 50.5
Abf.loc['Plastic (kg)', 'Tank (piece)'] =3.2

Abf.loc['Electricity (N. Gas kWh)', 'Powertrain (piece)'] = 2240
Abf.loc['N. Gas (MJ)', 'Powertrain (piece)'] = 2456
Abf.loc['Diesel (MJ)', 'Powertrain (piece)'] = 518
Abf.loc['Aluminium (kg)', 'Powertrain (piece)'] = 32
Abf.loc['Steel (kg)', 'Powertrain (piece)'] = 40
Abf.loc['Plastic (kg)', 'Powertrain (piece)'] = 2.3

Abf.loc['Steel (kg)', 'Gearbox (piece)'] = 24
Abf.loc['Aluminium (kg)', 'Gearbox (piece)'] = 4.6
Abf.loc['Electricity (N. Gas kWh)', 'Gearbox (piece)'] = 610

Abf = Abf.fillna(0)
display(Abf)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Electricity (N. Gas kWh)     100000.0        8110.0          50.5   
Steel (kg)                        0.0         300.0           0.0   
Aluminium (kg)                    0.0          70.0           0.0   
Platinum (kg)                     0.0           0.0           0.0   
Plastic (kg)                      0.0          65.0           3.2   
Diesel (MJ)                    2394.0         454.0           0.0   
N. Gas (MJ)                    9061.0        3110.0           0.0   
Transport (Truck, tkm)            0.0           0.0           0.0   
Concrete (kg)                     0.0           0.0           0.0   
Glass (kg)                       42.0           0.0           0.0   

                          Powertrain (piece)  Gearbox (piece)  
Electricity (N. Gas kWh)              2240.0            610.0  
Steel (kg)                              40.0             24.0  
Aluminium (kg)                          32.0              4.6  
Platinum (
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e0835797fd793596, points: 5, grade: true, solution: false -->
```python
assert Abf.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Abf.index == ['Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The index do not correspond'
assert (Abf.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)']).all(), 'The columns do not correspond'
```

#### **2c)** _(5 points)_ Fill in `Afb`

<!-- grade_id: cell-20c1e381dce990a8, points: 0, grade: false, solution: true -->
```python
Afb = Afb.fillna(0)
display(Afb)
```

Output:
```
Electricity (N. Gas kWh)  Steel (kg)  Aluminium (kg)  \
Car (piece)                              0.0         0.0             0.0   
Body (piece)                             0.0         0.0             0.0   
Tank (piece)                             0.0         0.0             0.0   
Powertrain (piece)                       0.0         0.0             0.0   
Gearbox (piece)                          0.0         0.0             0.0   

                    Platinum (kg)  Plastic (kg)  Diesel (MJ)  N. Gas (MJ)  \
Car (piece)                   0.0           0.0          0.0          0.0   
Body (piece)                  0.0           0.0          0.0          0.0   
Tank (piece)                  0.0           0.0          0.0          0.0   
Powertrain (piece)            0.0           0.0          0.0          0.0   
Gearbox (piece)               0.0           0.0          0.0          0.0   

                    Transport (Truck, tkm)  Concrete (kg)  Glass (kg)  
Car (piece)                 
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-d8dc30b5ec70f09b, points: 5, grade: true, solution: false -->
```python
assert Afb.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Afb.index == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)']).all(), 'The index do not correspond'
assert (Afb.columns == ['Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The columns do not correspond'
```

#### **2d)** _(5 points)_ Assemble `Aff`, `Abf`, `Afb` and `Abb` into `A` matrix

<!-- grade_id: cell-14596ab3b8ab17ae, points: 0, grade: false, solution: true -->
```python
A1 = pd.concat([Aff, Afb], axis = 1)
A2 = pd.concat([Abf, Abb], axis = 1)
A = pd.concat([A1, A2], axis = 0)
display(A)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Car (piece)                       0.0           0.0           0.0   
Body (piece)                      1.0           0.0           0.0   
Tank (piece)                      1.0           0.0           0.0   
Powertrain (piece)                1.0           0.0           0.0   
Gearbox (piece)                   1.0           0.0           0.0   
Electricity (N. Gas kWh)     100000.0        8110.0          50.5   
Steel (kg)                        0.0         300.0           0.0   
Aluminium (kg)                    0.0          70.0           0.0   
Platinum (kg)                     0.0           0.0           0.0   
Plastic (kg)                      0.0          65.0           3.2   
Diesel (MJ)                    2394.0         454.0           0.0   
N. Gas (MJ)                    9061.0        3110.0           0.0   
Transport (Truck, tkm)            0.0           0.0           0.0   
Concrete (kg)                     0.0           0.0         
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-76725c296f94c5e7, points: 5, grade: true, solution: false -->
```python
assert (A.index == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The index do not correspond'
assert (A.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The columns do not correspond'
```

### **S matrix** _(20 points)_

Run the cell below to import `Sf` and `Sb` as Pandas DataFrames. Do not change the names of the index or columns of the matrices nor their order as the autograding process might return an incorrect answer.

```python
Sf = pd.read_csv(files_folder + 'Sf.csv', sep = ',', index_col = [0])
Sb = pd.read_csv(files_folder + 'Sb.csv', sep = ',', index_col = [0])
```

#### **2e)** _(10 points)_ Fill in `Sf`

<!-- grade_id: cell-cc2ac2af5cb6d65e, points: 0, grade: false, solution: true -->
```python
Sf = Sf.fillna(0)
display(Sf)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Bauxite (kg)                     0.00         0.000        0.0000   
CH4 (kg)                         5.10         0.088        0.0071   
CO2 (kg)                      9420.00      1740.000      135.0000   
Crude Oil in Res. (kg)           0.00         0.000        0.0000   
Iron Ore (kg)                    0.00         0.000        0.0000   
Natural Gas in Res. (m3)         0.00         0.000        0.0000   
NOx (kg)                       156.30        38.600        0.6120   
Pt (kg)                          0.00         0.000        0.0000   
PM10 (kg)                        9.81         0.347        0.0031   
SO2 (kg)                        38.70        18.100        0.0920   

                          Powertrain (piece)  Gearbox (piece)  
Bauxite (kg)                           0.000            0.000  
CH4 (kg)                               0.120            0.028  
CO2 (kg)                            1480.000          889.000  
Crude Oil 
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-f6b5ad8ff815e44c, points: 10, grade: true, solution: false -->
```python
assert Sf.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Sf.index == ['Bauxite (kg)', 'CH4 (kg)', 'CO2 (kg)', 'Crude Oil in Res. (kg)',
                     'Iron Ore (kg)', 'Natural Gas in Res. (m3)', 'NOx (kg)', 'Pt (kg)',
                     'PM10 (kg)', 'SO2 (kg)']).all(), 'The index do not correspond'
assert (Sf.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)',
                       'Gearbox (piece)']).all(), 'The columns do not correspond'
```

#### **2f)** _(10 points)_ Assemble `Sf` and `Sb` into `S` matrix

<!-- grade_id: cell-64f7095bcd2a08d0, points: 0, grade: false, solution: true -->
```python
S = pd.concat([Sf, Sb], axis = 1)
display(S)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Bauxite (kg)                     0.00         0.000        0.0000   
CH4 (kg)                         5.10         0.088        0.0071   
CO2 (kg)                      9420.00      1740.000      135.0000   
Crude Oil in Res. (kg)           0.00         0.000        0.0000   
Iron Ore (kg)                    0.00         0.000        0.0000   
Natural Gas in Res. (m3)         0.00         0.000        0.0000   
NOx (kg)                       156.30        38.600        0.6120   
Pt (kg)                          0.00         0.000        0.0000   
PM10 (kg)                        9.81         0.347        0.0031   
SO2 (kg)                        38.70        18.100        0.0920   

                          Powertrain (piece)  Gearbox (piece)  \
Bauxite (kg)                           0.000            0.000   
CH4 (kg)                               0.120            0.028   
CO2 (kg)                            1480.000          889.000   
Crude 
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-63c6bbda9972d4eb, points: 10, grade: true, solution: false -->
```python
assert (S.index == ['Bauxite (kg)', 'CH4 (kg)', 'CO2 (kg)', 'Crude Oil in Res. (kg)',
                     'Iron Ore (kg)', 'Natural Gas in Res. (m3)', 'NOx (kg)', 'Pt (kg)',
                     'PM10 (kg)', 'SO2 (kg)']).all(), 'The index do not correspond'
assert (S.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The columns do not correspond'
```

## **3. Leontief inverse and output** _(20 points)_

Your task is to set up the calculate the Leontief inverse `L` and the output vector `x`.

#### **3a)** _(5 points)_ Preliminary steps

To calculate the Leontief inverse, an identity matrix `I` has be defined and the determinant of `(I-A)` has to be checked. Define the identity matrix `I` and calculate the determinant `detIA` of the matrix (I-A) in the cell under.

<!-- grade_id: cell-9955898d842a054d, points: 0, grade: false, solution: true -->
```python
I = np.identity(15)
detIA = np.linalg.det(I-A)
print(detIA)
```

Output:
```
0.981821240505319
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-975319ae2a3f6432, points: 5, grade: true, solution: false -->
```python
assert type(detIA) == np.float64, 'The answer(s) must be numeric'
```

#### **3b)** _(5 points)_ Compute the Leontief Inverse matrix `L`

Determine the L matrix that you will name `L`. Remember that `L` must be a Pandas DataFrame.

<!-- grade_id: cell-e86d00528228d70e, points: 0, grade: false, solution: true -->
```python
Larr = np.linalg.inv(I-A)
L = pd.DataFrame(Larr, index = ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)'], columns = ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)'])
display(L)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Car (piece)               1.000000e+00 -3.246896e-15 -6.479408e-18   
Body (piece)              1.000000e+00  1.000000e+00  1.420994e-16   
Tank (piece)              1.000000e+00  1.096332e-12  1.000000e+00   
Powertrain (piece)        1.000000e+00  5.332218e-14 -1.563933e-16   
Gearbox (piece)           1.000000e+00  5.842799e-13 -3.043125e-15   
Electricity (N. Gas kWh)  1.146289e+05  9.549494e+03  5.198797e+01   
Steel (kg)                4.675035e+02  3.090996e+02  5.266037e-02   
Aluminium (kg)            1.589907e+02  7.462956e+01  2.693941e-02   
Platinum (kg)             5.657249e-08  4.985088e-09  2.845353e-11   
Plastic (kg)              8.551389e+01  6.630434e+01  3.207281e+00   
Diesel (MJ)               2.012758e+04  1.932524e+03  8.485870e+00   
N. Gas (MJ)               1.315132e+06  1.116597e+05  5.909897e+02   
Transport (Truck, tkm)    6.771577e+03  5.974685e+02  3.428624e+00   
Concrete (kg)             6.059942e+01  7.16192
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-13b502370891cf4a, points: 5, grade: true, solution: false -->
```python
assert (L.index == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The index do not correspond'
assert (L.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)', 
                        'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The columns do not correspond'
```

#### **3c)** _(5 points)_ Fill in the final demand vector `y`

Run the cell below to import `y` as Pandas DataFrames. Do not change the names of the index or columns of the matrices nor their order as the autograding process might return an incorrect answer.

```python
y = pd.read_csv(files_folder + 'y.csv', sep = ',', index_col = [0])
```

Write in the cell under to fill in `y`

<!-- grade_id: cell-65fcc3b3619fbe25, points: 0, grade: false, solution: true -->
```python
display(y)
```

Output:
```
Final demand
Car (piece)                          1
Body (piece)                         0
Tank (piece)                         0
Powertrain (piece)                   0
Gearbox (piece)                      0
Electricity (N. Gas kWh)             0
Steel (kg)                           0
Aluminium (kg)                       0
Platinum (kg)                        0
Plastic (kg)                         0
Diesel (MJ)                          0
N. Gas (MJ)                          0
Transport (Truck, tkm)               0
Concrete (kg)                        0
Glass (kg)                           0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-309fc153b19ada31, points: 5, grade: true, solution: false -->
```python
assert y.isnull().values.any() == False, 'There are still NaN values in the vector'
assert (y.index == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)',
       'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The index do not correspond'
assert (y.columns == ['Final demand']).all(), 'The columns do not correspond'
```

#### **3d)** _(5 points)_ Compute the output vector `x` and rename its column into "Output"

<!-- grade_id: cell-fb73762396c2df0a, points: 0, grade: false, solution: true -->
```python
x = L @ y
x = x.rename(columns = {'Final demand' : 'Output'})
display(x)
```

Output:
```
Output
Car (piece)               1.000000e+00
Body (piece)              1.000000e+00
Tank (piece)              1.000000e+00
Powertrain (piece)        1.000000e+00
Gearbox (piece)           1.000000e+00
Electricity (N. Gas kWh)  1.146289e+05
Steel (kg)                4.675035e+02
Aluminium (kg)            1.589907e+02
Platinum (kg)             5.657249e-08
Plastic (kg)              8.551389e+01
Diesel (MJ)               2.012758e+04
N. Gas (MJ)               1.315132e+06
Transport (Truck, tkm)    6.771577e+03
Concrete (kg)             6.059942e+01
Glass (kg)                1.277542e+02
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-3d814cf86b24f721, points: 5, grade: true, solution: false -->
```python
assert (x.columns == ['Output']).all(), 'The column does not correspond'
```

## **4. Some interpretation questions** _(20 points)_

#### **4a)** _(10 points)_ Which one(s) of these statement(s) is/are correct?

1. The foreground processes denoted f come from a generic database.
2. There are 4 foreground processes in this system.
3. The foreground system is unidirectional.
 
Define your answer as the variable `answer_4a` (e.g. `answer_4a = {1}` if you think the correct statement is statement {1} and `answer_4a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-7a27ff1e1b88fb78, points: 0, grade: false, solution: true -->
```python
answer_4a = {3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-c5c3668e5f606cf6, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_4a]), 'The answer(s) must be numeric'
```

#### **4b)** _(10 points)_ Which one(s) of these statement(s) is/are correct?

1. The output vector x corresponds to the first column of the L matrix.
2. If we want a final demand of 2 vehicles, all the coefficients in the A matrix have to be multiplied by 2.
3. In the L matrix, we know how much electricity is needed to produce one unit output of steel.
 
Define your answer as the variable `answer_4b` (e.g. `answer_4b = {1}` if you think the correct statement is statement {1} and `answer_4b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-a6de82628b9f4dc1, points: 0, grade: false, solution: true -->
```python
answer_4b = {1, 2, 3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e3e50deebeec904e, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_4b]), 'The answer(s) must be numeric'
```

***

# **/!\ No help provided before the start of tutoring sessions for this part**

# **Problem Set 2 - Part b**

## **Passing requirements**

**All of the following must be satisfied for a passing grade**
- Correctly calculated e, E, d, Dpro and Dstr
- Correct interpretation of the matrices

## **5. Contribution analysis** _(70 points)_

#### **5a)** _(10 points)_ Compute the vector of total stressors `e` and rename its column into "Stressors"

<!-- grade_id: cell-1a3664f93dce4332, points: 0, grade: false, solution: true -->
```python
e = S @ x
e = e.rename(columns = {'Output' : 'Stressors'})
display(e)
```

Output:
```
Stressors
Bauxite (kg)              2.863966e+02
CH4 (kg)                  1.047526e+01
CO2 (kg)                  1.156845e+05
Crude Oil in Res. (kg)    3.523787e+04
Iron Ore (kg)             9.604991e+02
Natural Gas in Res. (m3)  3.538390e+05
NOx (kg)                  8.145112e+02
Pt (kg)                   1.174595e-07
PM10 (kg)                 1.287537e+01
SO2 (kg)                  7.763334e+01
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-bcca4983bdc25453, points: 10, grade: true, solution: false -->
```python
assert (e.columns == ['Stressors']).all(), 'The column does not correspond'
```

#### **5b)** _(10 points)_ Compute the matrix of stressors by processes `E`

<!-- grade_id: cell-ec7a9ffa1eea0a3b, points: 0, grade: false, solution: true -->
```python
x_i_marr = np.diag(x["Output"])
x_i_m = pd.DataFrame(x_i_marr, index = ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)',
       'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)'], columns = ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)',
       'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)'])
#display(x_i_m)
#display(S)

E = S @ x_i_m

display(E)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  \
Bauxite (kg)                     0.00         0.000        0.0000   
CH4 (kg)                         5.10         0.088        0.0071   
CO2 (kg)                      9420.00      1740.000      135.0000   
Crude Oil in Res. (kg)           0.00         0.000        0.0000   
Iron Ore (kg)                    0.00         0.000        0.0000   
Natural Gas in Res. (m3)         0.00         0.000        0.0000   
NOx (kg)                       156.30        38.600        0.6120   
Pt (kg)                          0.00         0.000        0.0000   
PM10 (kg)                        9.81         0.347        0.0031   
SO2 (kg)                        38.70        18.100        0.0920   

                          Powertrain (piece)  Gearbox (piece)  \
Bauxite (kg)                           0.000            0.000   
CH4 (kg)                               0.120            0.028   
CO2 (kg)                            1480.000          889.000   
Crude 
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-72489b4cff580cf9, points: 10, grade: true, solution: false -->
```python
assert (E.index == ['Bauxite (kg)', 'CH4 (kg)', 'CO2 (kg)', 'Crude Oil in Res. (kg)',
       'Iron Ore (kg)', 'Natural Gas in Res. (m3)', 'NOx (kg)', 'Pt (kg)',
       'PM10 (kg)', 'SO2 (kg)']).all(), 'The index do not correspond'
assert (E.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)',
       'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The columns do not correspond'
```

#### **5c)** _(10 points)_ Compute the vector of impacts `d` and rename its column into "Total impacts"

```python
C = pd.read_csv(files_folder + 'C.csv', sep = ',', index_col = [0])
```

<!-- grade_id: cell-cb2669c90cf7c4d2, points: 0, grade: false, solution: true -->
```python
d = C @ e
display(C)
d = d.rename(columns = {'Stressors' : 'Total impacts'})
display(d)
```

Output:
```
Bauxite (kg)  CH4 (kg)  CO2 (kg)  Crude Oil in Res. (kg)  Iron Ore (kg)  \
ADP  2.100000e-09         0         0                  0.0201   4.800000e-08   
GWP  0.000000e+00        23         1                  0.0000   0.000000e+00   
HTP  0.000000e+00         0         0                  0.0000   0.000000e+00   
AP   0.000000e+00         0         0                  0.0000   0.000000e+00   

     Natural Gas in Res. (m3)  NOx (kg)  Pt (kg)  PM10 (kg)  SO2 (kg)  
ADP                    0.0187       0.0     1.29       0.00     0.000  
GWP                    0.0000       0.0     0.00       0.00     0.000  
HTP                    0.0000       1.2     0.00       0.82     0.096  
AP                     0.0000       0.7     0.00       0.00     1.000  
     Total impacts
ADP    7325.071139
GWP  115925.399673
HTP     995.424074
AP      647.791197
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-0364a3816c7fe2f7, points: 10, grade: true, solution: false -->
```python
assert (d.columns == ['Total impacts']).all(), 'The column does not correspond'
```

#### **5d)** _(20 points)_ Compute the matrix of impacts by processes `Dpro`

<!-- grade_id: cell-78d3c5c59befc295, points: 0, grade: false, solution: true -->
```python
Dpro = C @ E
display(Dpro)
```

Output:
```
Car (piece)  Body (piece)  Tank (piece)  Powertrain (piece)  \
ADP       0.0000       0.00000      0.000000             0.00000   
GWP    9537.3000    1742.02400    135.163300          1482.76000   
HTP     199.3194      48.34214      0.745774            27.04024   
AP      148.1100      45.12000      0.520400            29.17000   

     Gearbox (piece)  Electricity (N. Gas kWh)  Steel (kg)  Aluminium (kg)  \
ADP         0.000000                713.824858    0.789593        4.841903   
GWP       889.644000              97987.967474  434.860116      281.082349   
HTP         5.171716                678.280233    0.707865        4.933195   
AP          3.021000                396.149350    0.471524        4.308648   

     Platinum (kg)  Plastic (kg)  Diesel (MJ)  N. Gas (MJ)  \
ADP       0.000001      0.372960   432.883870  5902.312550   
GWP       0.000085      8.136698  2272.158258     0.345354   
HTP       0.000010      0.004635    11.562127     5.718478   
AP        0.000014      0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-1ddebe50c492c119, points: 20, grade: true, solution: false -->
```python
assert (Dpro.index == ['ADP', 'GWP', 'HTP', 'AP']).all(), 'The index do not correspond'
assert (Dpro.columns == ['Car (piece)', 'Body (piece)', 'Tank (piece)', 'Powertrain (piece)',
       'Gearbox (piece)', 'Electricity (N. Gas kWh)', 'Steel (kg)', 'Aluminium (kg)',
       'Platinum (kg)', 'Plastic (kg)', 'Diesel (MJ)', 'N. Gas (MJ)',
       'Transport (Truck, tkm)', 'Concrete (kg)', 'Glass (kg)']).all(), 'The columns do not correspond'
```

#### **5e)** _(20 points)_ Compute the matrix of impacts by stressors `Dstr`

<!-- grade_id: cell-295476b638fb24c1, points: 0, grade: false, solution: true -->
```python
e_i_marr = np.diag(e["Stressors"])
e_i_m = pd.DataFrame(e_i_marr, index = ['Bauxite (kg)', 'CH4 (kg)', 'CO2 (kg)', 'Crude Oil in Res. (kg)',
       'Iron Ore (kg)', 'Natural Gas in Res. (m3)', 'NOx (kg)', 'Pt (kg)',
       'PM10 (kg)', 'SO2 (kg)'], columns = ['Bauxite (kg)', 'CH4 (kg)', 'CO2 (kg)', 'Crude Oil in Res. (kg)',
       'Iron Ore (kg)', 'Natural Gas in Res. (m3)', 'NOx (kg)', 'Pt (kg)',
       'PM10 (kg)', 'SO2 (kg)'])
display(e_i_m)

Dstr = C @ e_i_m
display(Dstr)
```

Output:
```
Bauxite (kg)  CH4 (kg)      CO2 (kg)  \
Bauxite (kg)                286.396646   0.00000       0.00000   
CH4 (kg)                      0.000000  10.47526       0.00000   
CO2 (kg)                      0.000000   0.00000  115684.46869   
Crude Oil in Res. (kg)        0.000000   0.00000       0.00000   
Iron Ore (kg)                 0.000000   0.00000       0.00000   
Natural Gas in Res. (m3)      0.000000   0.00000       0.00000   
NOx (kg)                      0.000000   0.00000       0.00000   
Pt (kg)                       0.000000   0.00000       0.00000   
PM10 (kg)                     0.000000   0.00000       0.00000   
SO2 (kg)                      0.000000   0.00000       0.00000   

                          Crude Oil in Res. (kg)  Iron Ore (kg)  \
Bauxite (kg)                             0.00000       0.000000   
CH4 (kg)                                 0.00000       0.000000   
CO2 (kg)                                 0.00000       0.000000   
Crude Oil in Res. (kg)         
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-0dd7e1d99995a01e, points: 20, grade: true, solution: false -->
```python
assert (Dstr.index == ['ADP', 'GWP', 'HTP', 'AP']).all(), 'The index do not correspond'
assert (Dstr.columns == ['Bauxite (kg)', 'CH4 (kg)', 'CO2 (kg)', 'Crude Oil in Res. (kg)',
       'Iron Ore (kg)', 'Natural Gas in Res. (m3)', 'NOx (kg)', 'Pt (kg)',
       'PM10 (kg)', 'SO2 (kg)']).all(), 'The columns do not correspond'
```

## **6. Some interpretation questions** _(30 points)_

#### **6a)** _(10 points)_ Which one(s) of these statement(s) is/are correct?

1. For the GWP impact category, 1 kg CO$_2$ contributes more than 1 kg CH$_4$.
2. For the HTP impact category, 1 kg NO$_x$ has a lower contribution than 1 kg SO$_2$.
3. For the AP impact category, 10 kg NO$_x$ has the same contribution than 7 kg SO$_2$.
 
Define your answer as the variable `answer_6a` (e.g. `answer_6a = {1}` if you think the correct statement is statement {1} and `answer_6a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-ceee56bf0b79296b, points: 0, grade: false, solution: true -->
```python
answer_6a = {3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-a552bfe9c37d31ef, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6a]), 'The answer(s) must be a set of values'
```

#### **6b)** _(10 points)_ Which one(s) of these statement(s) is/are correct?

1. Electricity gives the most contribution to ADP when 1 vehicle is produced.
2. Electricity gives the most contribution to GWP when 1 vehicle is produced.
3. Electricity gives the most contribution to HTP when 1 vehicle is produced.
4. Electricity gives the most contribution to AP when 1 vehicle is produced.
 
Define your answer as the variable `answer_6b` (e.g. `answer_6b = {1}` if you think the correct statement is statement {1} and `answer_6b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-b30b08ea8a75a192, points: 0, grade: false, solution: true -->
```python
answer_6b = {2,3,4}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e8b4903641b9b516, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6b]), 'The answer(s) must be a set of values'
```

#### **6c)** _(10 points)_ Which one(s) of these statement(s) is/are correct?

1. CO$_2$ does not contribute to ADP when 1 vehicle is produced.
2. CO$_2$ gives the most contribution to GWP when 1 vehicle is produced.
3. NO$_x$ gives the most contribution to HTP and AP when 1 vehicle is produced.
 
Define your answer as the variable `answer_6c` (e.g. `answer_6c = {1}` if you think the correct statement is statement {1} and `answer_6c = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-8f8af10b0468d688, points: 0, grade: false, solution: true -->
```python
answer_6c = {1,2, 3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-144d89a871acad4c, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_6c]), 'The answer(s) must be a set of values'
```