# **Problem Set 5**

## **Tasks**
- Set up Zff, Zfb and Zbf and the unadjusted Z matrix
- Calculate unadjusted x vector, foreground price vector and Hfb matrix
- Calculate modified Zbb matrix
- Calculate modified x vector
- Calculate A matrix
- Answering multiple choice questions

## **Passing requirement**
- **At least 80 points**

**_Please run the cell below before you start the assignment_**

```python
# Import required Python libraries
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns

# Folder for the files
files_folder = "files_PS5_V1_0/"

### In case formulas are entered in the csv files & correcting the datatypes
def convert_formulas(matrix):
    for row in matrix.index:
        for column in matrix.columns:
            if type(matrix.loc[row, column]) == str:
                if " " in matrix.loc[row, column]:
                    matrix.loc[row, column] = matrix.loc[row, column].lstrip(' ')
                if "=" in matrix.loc[row, column]:
                    matrix.loc[row, column] = matrix.loc[row, column].lstrip('=')
                matrix.loc[row, column] = pd.eval(matrix.loc[row, column])
    ### Get the data types right in the DataFrames
    s = matrix.select_dtypes(include='object').columns
    matrix[s] = matrix[s].astype("float64")
    return matrix
```

## **Integrated hybrid LCA**

The Norwegian government is interested in analysing the prospect of ethanol production from bioresources, to be used as biofuel. The main current alcohol production in Norway is the production of aquavit from potatoes. Hence this ethanol production will be the case study. To analyse this, the crops must be disaggregated from the agricultural sector and the ethanol distillation must be disaggregated from the chemicals, oil and gas sector. It is your task to calculate the requirements matrix with Norway’s potato production disaggregated from the agricultural sector and aquavit distillation disaggregated from the chemicals, oil and gas sector.

You are given a background system, which represents a simplified version of the flows in the Norwegian economy. The background system has four sectors: Chemicals, oil and gas sector; Food manufacturing sector; Agricultural sector; and all other sectors aggregated into one sector called Other. The background system’s flow matrix `Zbb` is as follows:


|Zbb, in billion NOK|||||
|-|-|-|-|-|
||**Chemicals, oil and gas**|**Food manufacturing**|**Agriculture**|**Other**|
|**Chemicals, oil and gas**|10.4|7.6|0.008|60.6|
|**Food manufacturing**|3.5|38|5|30.6|
|**Agriculture**|0.37|21.6|1.5|3.2|
|**Other**|74|58|8.6|1600|

And the total output from the background `xb` is as follows:

|xb, in billion NOK||
|-|-|
|**Chemicals, oil and gas**|700|
|**Food manufacturing**|170|
|**Agriculture**|35|
|**Other**|3114|

You shall now disaggregate the foreground processes potato production and aquavit distillation from the flows in the Norwegian economy. We know the following:
- The national aquavit distillation uses 40 000 tons potatoes per year, 140 million NOK worth of inputs from the food manufacturing sector and 110 million NOK from the other sectors, to produce a total output of 5500 tons of aquavit.
- The potato production requires 350 000 NOK per year from the oil sector, 1 billion NOK from the food manufacturing sector, 225 million NOK from the agricultural sector and 1.2 billion from the other sectors; to produce a total output of 350 000 tons of potatoes.
- The food manufacturing sector uses 150 000 tons of potatoes.
- The agriculture sector uses 45 000 tons potatoes as feed (assume 0 to the potato production).
- All of the aquavit production and the potato output that does not go into aquavit production, agriculture feed or food manufacturing, is assumed to go to final demand.
- The price of potatoes is 2000 NOK/ton, and the price of aquavit is 450 000 NOK/ton.

## **1. The unadjusted `Z` matrix** _(25 points)_

Your task is to set up the unadjusted `Z` matrix. `Zbb` is already given. You have to fill in the matrices `Zff`, `Zbf`, `Zfb`. For that, you have two options: 
- Option 1: filling in the matrices in the csv files located in the **files_PS5_V1_0** folder and save the csv files by using "ctrl+s". **If you are filling in the matrices that way and are calculating some numbers, you should enter the formulas in the table.**
- Option 2: filling in the matrices in the cells in the sections 1a), 1b), 1c). If you go for that option, we recommend you to make a quick hand sketch of the matrices to visualize how they should look like.

```python
Zff = convert_formulas(pd.read_csv(files_folder + 'Zff.csv', sep = ',', index_col = [0]))
Zfb = convert_formulas(pd.read_csv(files_folder + 'Zfb.csv', sep = ',', index_col = [0]))
Zbf = convert_formulas(pd.read_csv(files_folder + 'Zbf.csv', sep = ',', index_col = [0]))
Zbb = convert_formulas(pd.read_csv(files_folder + 'Zbb.csv', sep = ',', index_col = [0]))
```

#### **1a)** _(5 points)_ Fill in `Zff`

<!-- grade_id: cell-92c639dfbf7e0f99, points: 0, grade: false, solution: true -->
```python
Zff = Zff.fillna(0)
display(Zff)
```

Output:
```
Aquavit distillation  Potato
Aquavit distillation                   0.0     0.0
Potato                             40000.0     0.0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-952b3b7439df92f2, points: 5, grade: true, solution: false -->
```python
assert Zff.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Zff.index == ['Aquavit distillation', 'Potato']).all(), 'The index do not correspond'
assert (Zff.columns == ['Aquavit distillation', 'Potato']).all(), 'The columns do not correspond'
```

#### **1b)** _(5 points)_ Fill in `Zfb`

<!-- grade_id: cell-0fb2ccf13f50bf3b, points: 0, grade: false, solution: true -->
```python
Zfb = Zfb.fillna(0)
display(Zfb)
```

Output:
```
Chemicals, oil and gas  Food manufacturing  Agriculture  \
Aquavit distillation                     0.0                 0.0          0.0   
Potato                                   0.0            150000.0      45000.0   

                      Other  
Aquavit distillation    0.0  
Potato                  0.0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-811f19a0bf13c2d2, points: 5, grade: true, solution: false -->
```python
assert Zfb.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Zfb.index == ['Aquavit distillation', 'Potato']).all(), 'The index do not correspond'
assert (Zfb.columns == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The columns do not correspond'
```

#### **1c)** _(10 points)_ Fill in `Zbf`

<!-- grade_id: cell-fcc40d20b6e67896, points: 0, grade: false, solution: true -->
```python
Zbf = Zbf.fillna(0)
display(Zbf)
```

Output:
```
Aquavit distillation   Potato
Chemicals, oil and gas                  0.00  0.00035
Food manufacturing                      0.14  1.00000
Agriculture                             0.00  0.22500
Other                                   0.11  1.20000
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-1b97de89d8303de2, points: 10, grade: true, solution: false -->
```python
assert Zbf.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Zbf.index == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (Zbf.columns == ['Aquavit distillation', 'Potato']).all(), 'The columns do not correspond'
```

#### **1d)** _(5 points)_ Assemble `Zff`, `Zbf`, `Zfb` and `Zbb` into `Z` matrix

<!-- grade_id: cell-4d30564641748d62, points: 0, grade: false, solution: true -->
```python
Z1 = pd.concat([Zff, Zfb], axis = 1)
Z2 = pd.concat([Zbf, Zbb], axis = 1)
Z = pd.concat([Z1, Z2], axis = 0)
display(Z)
```

Output:
```
Aquavit distillation   Potato  Chemicals, oil and gas  \
Aquavit distillation                    0.00  0.00000                    0.00   
Potato                              40000.00  0.00000                    0.00   
Chemicals, oil and gas                  0.00  0.00035                   10.40   
Food manufacturing                      0.14  1.00000                    3.50   
Agriculture                             0.00  0.22500                    0.37   
Other                                   0.11  1.20000                   74.00   

                        Food manufacturing  Agriculture   Other  
Aquavit distillation                   0.0        0.000     0.0  
Potato                            150000.0    45000.000     0.0  
Chemicals, oil and gas                 7.6        0.008    60.6  
Food manufacturing                    38.0        5.000    30.6  
Agriculture                           21.6        1.500     3.2  
Other                                 58.0        8.600  160
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-6824cb5fe8a1e1fe, points: 5, grade: true, solution: false -->
```python
assert (Z.index == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas',
       'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (Z.columns == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas',
       'Food manufacturing', 'Agriculture', 'Other']).all(), 'The columns do not correspond'
```

## **2. The unadjusted `x` vector** _(15 points)_

Your task is to set up the unadjusted `X` vector. You have to fill in `xf`, `xb`. For that, you have two options: 
- Option 1: filling in the vectors in the csv files located in the **files_PS5_V1_0** folder and save the csv files by using "ctrl+s". **If you are filling in the vectors that way and are calculating some numbers, you should enter the formulas in the table.**
- Option 2: filling in the vectors in the cells in the sections 2a), 2b). If you go for that option, we recommend you to make a quick hand sketch of the vectors to visualize how they should look like.

```python
xf = convert_formulas(pd.read_csv(files_folder + 'xf.csv', sep = ',', index_col = [0]))
xb = convert_formulas(pd.read_csv(files_folder + 'xb.csv', sep = ',', index_col = [0]))
```

#### **2a)** _(5 points)_ Fill in `xf`

<!-- grade_id: cell-3573f7f07a8aaf23, points: 0, grade: false, solution: true -->
```python
xf = xf.fillna(0)
display(xf)
```

Output:
```
Output
Aquavit distillation    5500
Potato                350000
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e835e128e0e7ab53, points: 5, grade: true, solution: false -->
```python
assert xf.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (xf.index == ['Aquavit distillation', 'Potato']).all(), 'The index do not correspond'
assert (xf.columns == ['Output']).all(), 'The columns do not correspond'
```

#### **2b)** _(5 points)_ Fill in `xb`

<!-- grade_id: cell-37e3f5827fd355aa, points: 0, grade: false, solution: true -->
```python
xb = xb.fillna(0)
display(xb)
```

Output:
```
Output
Chemicals, oil and gas     700
Food manufacturing         170
Agriculture                 35
Other                     3114
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e9ee7df0b26631f4, points: 5, grade: true, solution: false -->
```python
assert xb.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (xb.index == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (xb.columns == ['Output']).all(), 'The columns do not correspond'
```

#### **2c)** _(5 points)_ Assemble `xf` and `xb` into `x` vector

<!-- grade_id: cell-57a13c41d45ab49d, points: 0, grade: false, solution: true -->
```python
x = pd.concat([xf, xb], axis = 0)
display(x)
```

Output:
```
Output
Aquavit distillation      5500
Potato                  350000
Chemicals, oil and gas     700
Food manufacturing         170
Agriculture                 35
Other                     3114
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-37725338ea1b6942, points: 5, grade: true, solution: false -->
```python
assert (x.index == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas',
       'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (x.columns == ['Output']).all(), 'The columns do not correspond'
```

## **3. The foreground price vector `pi`** _(5 points)_

Your task is to set up the foreground price vector `pi`. You have to fill in `pi`. For that, you have two options: 
- Option 1: filling in the vector in the csv files located in the **files_PS5_V1_0** folder and save the csv files by using "ctrl+s". **If you are filling in the vector that way and are calculating some numbers, you should enter the formulas in the table.**
- Option 2: filling in the vector in the cell under. If you go for that option, we recommend you to make a quick hand sketch of the vector to visualize how it should look like.

```python
pi = convert_formulas(pd.read_csv(files_folder + 'pi.csv', sep = ',', index_col = [0]))
```

Fill in `pi`

<!-- grade_id: cell-f86e80e840d1cbe3, points: 0, grade: false, solution: true -->
```python
pi = pi.fillna(0)
display(pi)
```

Output:
```
Unit cost
Aquavit distillation   0.000450
Potato                 0.000002
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-0e04d1ebc1a72cc1, points: 5, grade: true, solution: false -->
```python
assert pi.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (pi.index == ['Aquavit distillation', 'Potato']).all(), 'The index do not correspond'
assert (pi.columns == ['Unit cost']).all(), 'The columns do not correspond'
```

## **4. Foreground-background correspondence matrix `Hfb`** _(5 points)_

Your task is to set up the foreground-background correspondence matrix `Hfb`. You have to fill in `Hfb`. For that, you have two options: 
- Option 1: filling in the matrix in the csv files located in the **files_PS5_V1_0** folder and save the csv files by using "ctrl+s". **If you are filling in the matrix that way and are calculating some numbers, you should enter the formulas in the table.**
- Option 2: filling in the matrix in the cell under. If you go for that option, we recommend you to make a quick hand sketch of the matrix to visualize how it should look like.

```python
Hfb = convert_formulas(pd.read_csv(files_folder + 'Hfb.csv', sep = ',', index_col = [0]))
```

Fill in `Hfb`

<!-- grade_id: cell-7f9a42634955bd30, points: 0, grade: false, solution: true -->
```python
Hfb = Hfb.fillna(0)
display(Hfb)
```

Output:
```
Chemicals, oil and gas  Food manufacturing  Agriculture  \
Aquavit distillation                     1.0                 0.0          0.0   
Potato                                   0.0                 0.0          1.0   

                      Other  
Aquavit distillation    0.0  
Potato                  0.0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-e1e35b1a0ef2f0cb, points: 5, grade: true, solution: false -->
```python
assert Hfb.isnull().values.any() == False, 'There are still NaN values in the matrix'
assert (Hfb.index == ['Aquavit distillation', 'Potato']).all(), 'The index do not correspond'
assert (Hfb.columns == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The columns do not correspond'
```

## **5. Calculate the modified `Zbb` and `Z` matrices: `Zbb_adj` and `Z_adj`** _(20 points)_

#### **5a)** _(15 points)_ Calculate `Zbb_adj`

<!-- grade_id: cell-3423911ea8f4b8a4, points: 0, grade: false, solution: true -->
```python
Hfb_transpose = np.transpose(Hfb)
#display(Hfb_transpose)

pi_arr = np.diag(pi["Unit cost"])
pi_hat = pd.DataFrame(pi_arr, index = ['Aquavit distillation', 'Potato'], columns = ['Aquavit distillation', 'Potato'])
#display(pi_hat)

Zbf_adj = Zbf @ Hfb
#display(Zbf_adj)

Zfb_adj = Hfb_transpose @ pi_hat @ Zfb
#display(Zfb_adj)

Zff_adj = Hfb_transpose @ pi_hat @ Zff @ Hfb
#display(Zff_adj)

Zbb_adj = Zbb - Zff_adj - Zfb_adj -Zbf_adj
display(Zbb_adj)
```

Output:
```
Chemicals, oil and gas  Food manufacturing  \
Chemicals, oil and gas                   10.40                 7.6   
Food manufacturing                        3.36                38.0   
Agriculture                               0.29                21.3   
Other                                    73.89                58.0   

                        Agriculture   Other  
Chemicals, oil and gas      0.00765    60.6  
Food manufacturing          4.00000    30.6  
Agriculture                 1.18500     3.2  
Other                       7.40000  1600.0
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-c475b277025003a4, points: 15, grade: true, solution: false -->
```python
assert (Zbb_adj.index == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (Zbb_adj.columns == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The columns do not correspond'
```

#### **5b)** _(5 points)_ Calculate `Z_adj`

<!-- grade_id: cell-917c8bd7009cfc4d, points: 0, grade: false, solution: true -->
```python
Z1 = pd.concat([Zff, Zfb], axis = 1)
Z2 = pd.concat([Zbf, Zbb_adj], axis = 1)
Z_adj = pd.concat([Z1, Z2], axis = 0)
display(Z_adj)
```

Output:
```
Aquavit distillation   Potato  Chemicals, oil and gas  \
Aquavit distillation                    0.00  0.00000                    0.00   
Potato                              40000.00  0.00000                    0.00   
Chemicals, oil and gas                  0.00  0.00035                   10.40   
Food manufacturing                      0.14  1.00000                    3.36   
Agriculture                             0.00  0.22500                    0.29   
Other                                   0.11  1.20000                   73.89   

                        Food manufacturing  Agriculture   Other  
Aquavit distillation                   0.0      0.00000     0.0  
Potato                            150000.0  45000.00000     0.0  
Chemicals, oil and gas                 7.6      0.00765    60.6  
Food manufacturing                    38.0      4.00000    30.6  
Agriculture                           21.3      1.18500     3.2  
Other                                 58.0      7.40000  160
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-2af3b6f725d49056, points: 5, grade: true, solution: false -->
```python
assert (Z_adj.index == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (Z_adj.columns == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The columns do not correspond'
```

## **6. Calculate the modified `xb` and `x` vectors: `Xb_adj` and `X_adj`** _(10 points)_

#### **6a)** _(5 points)_ Calculate `Xb_adj`

<!-- grade_id: cell-7f4f9a229ad42b88, points: 0, grade: false, solution: true -->
```python
Xb_adj = xb - Hfb_transpose @ pi_hat @ xf
display(Xb_adj)
```

Output:
```
Output
Chemicals, oil and gas   697.525
Food manufacturing       170.000
Agriculture               34.300
Other                   3114.000
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-4b580b2c1ef20581, points: 5, grade: true, solution: false -->
```python
assert (Xb_adj.index == ['Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (Xb_adj.columns == ['Output']).all(), 'The columns do not correspond'
```

#### **6b)** _(5 points)_ Calculate `X_adj`

<!-- grade_id: cell-6a02ae2025c76641, points: 0, grade: false, solution: true -->
```python
X_adj = pd.concat([xf, Xb_adj], axis = 0)
display(X_adj)
```

Output:
```
Output
Aquavit distillation      5500.000
Potato                  350000.000
Chemicals, oil and gas     697.525
Food manufacturing         170.000
Agriculture                 34.300
Other                     3114.000
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-9cb062aa8783e32c, points: 5, grade: true, solution: false -->
```python
assert (X_adj.index == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (X_adj.columns == ['Output']).all(), 'The columns do not correspond'
```

## **7. Calculate the A matrix for this system** _(10 points)_

<!-- grade_id: cell-ba49f9e67bfb0d97, points: 0, grade: false, solution: true -->
```python
X_arr = np.diag(X_adj["Output"])
X_hat_inv_arr = np.linalg.inv(X_arr)
X_hat_inv = pd.DataFrame(X_hat_inv_arr, index = ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other'], columns = ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other'])
#display(X_hat_inv)

A = Z_adj @ X_hat_inv
display(A)
```

Output:
```
Aquavit distillation        Potato  \
Aquavit distillation                0.000000  0.000000e+00   
Potato                              7.272727  0.000000e+00   
Chemicals, oil and gas              0.000000  1.000000e-09   
Food manufacturing                  0.000025  2.857143e-06   
Agriculture                         0.000000  6.428571e-07   
Other                               0.000020  3.428571e-06   

                        Chemicals, oil and gas  Food manufacturing  \
Aquavit distillation                  0.000000            0.000000   
Potato                                0.000000          882.352941   
Chemicals, oil and gas                0.014910            0.044706   
Food manufacturing                    0.004817            0.223529   
Agriculture                           0.000416            0.125294   
Other                                 0.105932            0.341176   

                        Agriculture     Other  
Aquavit distillation       0.000000  0.000000  
Po
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-35a27735647e92cd, points: 10, grade: true, solution: false -->
```python
assert (A.index == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The index do not correspond'
assert (A.columns == ['Aquavit distillation', 'Potato', 'Chemicals, oil and gas', 'Food manufacturing', 'Agriculture', 'Other']).all(), 'The columns do not correspond'
```

## **8. Some interpretation questions** _(10 points)_

#### **8a)** _(5 points)_ Which one(s) of these statement(s) is/are correct?

By looking at `Zbb_adj`:
1. All the values of the initial `Zbb` matrix have been modified.
2. Only some values of the initial `Zbb` matrix have been modified.
3. The modified values are lower than in the initial `Zbb` matrix.
4. The modified values are higher than in the initial `Zbb` matrix.
 
Define your answer as the variable `answer_8a` (e.g. `answer_8a = {1}` if you think the correct statement is statement {1} and `answer_8a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-ceee56bf0b79296b, points: 0, grade: false, solution: true -->
```python
answer_8a = {2, 3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-a552bfe9c37d31ef, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_8a]), 'The answer(s) must be a set of values'
```

#### **8b)** _(5 points)_ Which one(s) of these statement(s) is/are correct?

By looking at `A`:
1. There are flows from the foreground to the background.
2. All the sectors/processes have the unit tons/ton.
3. We can get how many tons potato are required to produce 1 ton aquavit.
 
Define your answer as the variable `answer_8b` (e.g. `answer_8b = {1}` if you think the correct statement is statement {1} and `answer_8b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-98b8c1c6d2dedb59, points: 0, grade: false, solution: true -->
```python
answer_8b = {1, 3}
```

**_Run the cell under to check the format of your answer(s)_**

<!-- grade_id: cell-394302e0ccf69795, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_8b]), 'The answer(s) must be a set of values'
```