# **Problem Set 6**

## **Problem set description**

This exercise focuses on the phases of LCA, functional unit, regionalization, and understanding of DALYs.

## **Tasks**
- The 4 LCA phases and which question(s) each phase answers
- The definition of functional unit and its application
- Understanding of regionalization
- What DALY is and how it is calculated

## **Passing requirement**
- **At least 80 points**
- **Question 3b) has to be delivered on Blackboard in Assignments -> PS6 -> "PS6 - Questions BB"**
- **Question 4c) has to be delivered on Blackboard in Assignments -> PS6 -> "PS6 - Questions BB"**

**Answers to questions 3b) and 4c) count for 25 points in total. It is part of the 80 points needed for getting an approved assignment. Write your answers in the word template provided in Assignments -> PS6 -> "PS6 - Questions BB" converted to a pdf file!**

**_Please run the cell below before you start the assignment_**

```python
# Import required Python libraries
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
```

## **1. Phases of LCA** _(25 points)_

#### **1a)** LCA diagram _(5 points)_ 

In LCA, there are 4 phases: 
1) Goal and scope 
2) Inventory analysis
3) Impact assessment 
4) Interpretation

Place the names of the phases in the diagram below. Define your answer as the variables `answer_1a_A`, `answer_1a_B` and `answer_1a_C`. Your answer should be `answer_1a_A = {1}` if you think the correct name for A is "1) Goal and scope". The 5 points are only given if all the names are correctly placed!

<img src="files_PS6_V1_0/LCA_phases.svg" width="700">

<!-- grade_id: cell-bdbe7a31da54d78e, points: 0, grade: false, solution: true -->
```python
answer_1a_A = {2}
answer_1a_B = {4}
answer_1a_C = {3}
```

**_Run the cell under to check the format of your answers_**

<!-- grade_id: cell-92ea1755d5a3fe5f, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_1a_A, answer_1a_B, answer_1a_C]), 'The answer(s) must be a set of values'
```

#### **1b)** Which question(s) does each phase answer? _(20 points)_ 

You'll find a list of questions below, each associated with one of the four phases of the Life Cycle Assessment (LCA) mentioned in 1a). Your task is to match these LCA phases with the corresponding questions. Use the variables `answer_1b_1`, `answer_1b_2`, `answer_1b_3`, and `answer_1b_4` to define your responses for each phase, using letters a through m to represent the questions. 

For example, if you believe that only question 'a' is addressed in phase "1) Goal and scope," set `answer_1b_1 = {'a'}`. If you think both questions 'a', 'j' and 'l'  are addressed in phase 1, set it as `answer_1b_1 = {'a', 'j', 'l'}`. There is only one correct phase for each question, but there are several questions for each particular phase.

Each correct response will earn you 5 points. However, to receive these 5 points, your assignment of questions to a particular phase must be entirely accurate. Therefore, please make sure you match each question to the appropriate phase of the LCA. 


+ a - What is the aim of the LCA?
+ b - What are the possible recommendations?
+ c - Where are the system boundaries and what are the included functions in the system? 
+ d - How do the results look like after weighting/normalizing?
+ e - What are the weaknesses of the study? 
+ f - Which raw materials/emissions are used/produced? 
+ g - What assumptions are made? 
+ h - How reliable are the results? 
+ i - Which products are produced? Which products and materials are used? 
+ j - What are the results for different impact categories?
+ k - How are these different inputs and outputs allocated to the system? 
+ l - Conclusions? 

**_Questions from the list above about 1) Goal and scope `(answer_1b_1)`_**

<!-- grade_id: cell-435fd41fbde03ab4, points: 0, grade: false, solution: true -->
```python
answer_1b_1 = {'a', 'c', 'g'}
```

**_Run the cell under to check the format of your answers_**

<!-- grade_id: cell-98048c7111c3cd99, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_1b_1]), 'The answer(s) must be a set of values'
```

**_Questions from the list above about 2) Inventory analysis `(answer_1b_2)`_**

<!-- grade_id: cell-06482e56a29813e0, points: 0, grade: false, solution: true -->
```python
answer_1b_2 = {'f', 'i', 'k'}
```

**_Run the cell under to check the format of your answers_**

<!-- grade_id: cell-1018a6e0f68c8aab, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_1b_2]), 'The answer(s) must be a set of values'
```

**_Questions from the list above about 3) Impact assessment `(answer_1b_3)`_**

<!-- grade_id: cell-d2deb2f53c3057c6, points: 0, grade: false, solution: true -->
```python
answer_1b_3 = {'d', 'j'}
```

**_Run the cell under to check the format of your answers_**

<!-- grade_id: cell-3905f651a129466d, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_1b_3]), 'The answer(s) must be a set of values'
```

**_Questions from the list above about 4) Interpretation `(answer_1b_4)`_**

<!-- grade_id: cell-83b0fb99634f3a07, points: 0, grade: false, solution: true -->
```python
answer_1b_4 = {'b', 'e', 'h', 'l'}
```

**_Run the cell under to check the format of your answers_**

<!-- grade_id: cell-4be7a382e684f419, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_1b_4]), 'The answer(s) must be a set of values'
```

## **2. Functional units** _(30 points)_

#### **2a)** Which one(s) of the following statements correctly describe the concept of a functional unit in LCA? _(10 points)_ 

1. A functional unit is a quantified performance of a product system for use as a reference unit.
2. The functional unit is used to ensure that comparisons between different product systems are meaningful and consistent.
3. A functional unit is the weight of a product, which is used to compare different products.
4. The functional unit represents the specific function or service provided by a product or system, allowing for a fair comparison between alternatives.
5. A functional unit is the geographical location of a product's production, used for comparing environmental impacts.

Define your answer as the variable `answer_2a` (e.g. `answer_2a = {1}` if you think the correct statement is statement {1} or `answer_2a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-50e46c208f57f4b4, points: 0, grade: false, solution: true -->
```python
answer_2a = {1, 2, 4}
```

<!-- grade_id: cell-ebb17b2aa0bbb8eb, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_2a]), 'The answer(s) must be a set of values'
```

#### **2b)** Define appropriate functional units for comparing a remote control with non-rechargeable batteries and a remote control with rechargeable batteries _(10 points)_ 

1. Distance in meters between the remote control and the device it controls.
2. Cost of the remote control.
3. Total hours of remote control usage over its lifetime.
4. Number of button presses provided over the lifetime of the remote control.
5. Total weight of the remote control with batteries.

Define your answer as the variable `answer_2b` (e.g. `answer_2b = {1}` if you think the correct statement is statement {1} or `answer_2b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-b99a717b1bc8f233, points: 0, grade: false, solution: true -->
```python
answer_2b = {3, 4}
```

<!-- grade_id: cell-d5ce5db1baa9bf8f, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_2b]), 'The answer(s) must be a set of values'
```

#### **2c)** Define appropriate functional units for comparing an online encyclopedia and a printed encyclopaedia _(10 points)_ 

1. Total weight of the printed encyclopedia or the electronic device used to access the online encyclopedia.
2. Total number of articles read or accessed by a user over a defined period (e.g., one year).
3. Cost of the printed encyclopedia or subscription fee for the online version.
4. Number of colors used in the printed encyclopedia or displayed on the screen for the online version.
5. Total usage hours of users i one year.

Define your answer as the variable `answer_2c` (e.g. `answer_2c = {1}` if you think the correct statement is statement {1} or `answer_2c = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-d0af6f949bc9a5b6, points: 0, grade: false, solution: true -->
```python
answer_2c = {2, 5}
```

<!-- grade_id: cell-8ed82bf5b95c1884, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_2c]), 'The answer(s) must be a set of values'
```

## **3. Regionalization** _(25 points)_

#### **3a)** Which of the following are reasons why regionalization is attempted in many different impact categories in LCA?  _(10 points)_ 

1. Due to trade and global value chain, products are produced, used, and disposed of in different world regions, requiring various inputs and generating diverse outputs.
2. Some of the impacts are highly location-specific, and identifying these with a regionalized assessment allows for better understanding of supply chains and key processes in certain world regions.
3. Regionalization simplifies the assessment process by allowing the use of a single global average value for all impact categories.
4. Regionalized assessments of highly location-specific impacts help reduce the uncertainty of the assessments.
5. Regionalization is mainly used for marketing purposes to emphasize the environmental benefits of a product in a specific region.

Define your answer as the variable `answer_3a` (e.g. `answer_3a = {1}` if you think the correct statement is statement {1} or `answer_3a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-9aefb05e5674291f, points: 0, grade: false, solution: true -->
```python
answer_3a = {1, 2, 4}
```

<!-- grade_id: cell-755f22dc4fc5ce49, points: 10, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_3a]), 'The answer(s) must be a set of values'
```

#### **3b)** Explain the importance of regionalization by using one example of your choice _(15 points)_ 

You have to **deliver on Blackboard** a short paragraph (around 80 words is the recommended length) to give an example explaining the importance of regionalization.

## **4. Explain the concept of DALYs** _(20 points)_

#### **4a)** What is DALYs? _(5 points)_

1. A measure of the total number of years lived with a disability or injury.
2. The number of years lost due to ill-health, disability, or early death within a population.
3. The average number of years a person can expect to live in good health.
4. The total number of years a person is expected to live with a particular disease or disability.

Define your answer as the variable `answer_4a` (e.g. `answer_4a = {1}` if you think the correct statement is statement {1} or `answer_4a = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-2e633d867d3da27d, points: 0, grade: false, solution: true -->
```python
answer_4a = {2}
```

<!-- grade_id: cell-17aae40f31747b89, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_4a]), 'The answer(s) must be a set of values'
```

#### **4b)** Which of the following statements is NOT correct about DALYs? _(5 points)_

1. DALYs measure the number of years lost due to ill-health, disability, or early death. 
2. DALYs are used to quantify the overall disease burden.
3. DALYs combine the years of life lost due to premature mortality with the years of healthy life lost due to disability.
4. DALYs focus on the economic impacts of diseases and ignore the social and environmental aspects of health.

Define your answer as the variable `answer_4b` (e.g. `answer_4b = {1}` if you think the correct statement is statement {1} or `answer_4b = {1, 2}` if you think the correct statements are statements {1, 2}).

<!-- grade_id: cell-b2613a62b0fc32fd, points: 0, grade: false, solution: true -->
```python
answer_4b = {4}
```

<!-- grade_id: cell-602e14721ea16e3a, points: 5, grade: true, solution: false -->
```python
assert all (type(i) in [set] for i in [answer_4b]), 'The answer(s) must be a set of values'
```

#### **4c)** How is DALYs calculated? _(10 points)_

You have to **deliver on Blackboard** the formula with a description of all the variables involved in the formula. For the formula, make it easy to read by using the equation tool from Word for example.